# Cuadrados perfectos

Los cuadrados perfectos son los números $1$, $4$, $9$, 
$16$, ..., esto es, los cuadrados de los números naturales:
$1^2$, $2^2$, $3^2$, $4^2$,.....

Escribe una función que nos indique si un número natural $m$ es un cuadrado perfecto, o sea, si es de la forma $m = n^2$ para algún natural $n$.

<!-- grade_id: cell-6ea5308918f6bf42, points: 0, grade: false, solution: true -->
```python
import math

def perfect_square(n: int) -> bool:
    ### BEGIN SOLUTION
    return n>=0 and math.floor(math.sqrt(n))**2 == n
    ### END SOLUTION

```

```python
#Escribe aquí tus pruebas

```

<!-- grade_id: cell-edd0d845210286c4, points: 1, grade: true, solution: false -->
```python
def test_perfect_square():
    tests = [
        (-4, False),
        (-1, False),
        (0, True),
        (1, True),
        (2, False),
        (5, False),
    ]
    for num, res in tests:
        r = perfect_square(num)
        print(f'{num}, {r}.....', end='')
        assert res == r, f'Error, debería ser {res}'
        print('OK')
        
    for i in range(0, 100, 7):
        assert perfect_square(i*i), f'Error {i*i} es un cuadrado perfecto'
    print('OK')
        

test_perfect_square()
```

Output:
```
-4, False.....OK
-1, False.....OK
0, True.....OK
1, True.....OK
2, False.....OK
5, False.....OK
OK
```

Escribe una función que nos devuelva el número $n$ mayor posible pero que no supere a $m$,
(o sea, $n\leq m$) y sea cuadrado perfecto.

<!-- grade_id: cell-a94a064a5589ebaf, points: 0, grade: false, solution: true -->
```python
def previous_square(n: int) -> float:
    ### BEGIN SOLUTION
    assert n>=0, f'Error no se puede calcular la raíz de un número negativo:{n} '
    return math.floor(math.sqrt(n))**2
    ### END SOLUTION

```

<!-- grade_id: cell-9853f8979e61aaed, points: 1, grade: true, solution: false -->
```python
### Escribe aquí tus pruebas

```

```python
def test_previous_square():
    for i in range(100):
        for j in range(i**2, i**2 + 1):
            assert previous_square(j)==i**2, \
                f'El cuadrado anterior a {j} es {i**2}'

    for i in range(-1, -100, -1):
        try:
            n = previous_square(-5)
        except AssertionError:
            pass
        except:
            assert False, f'La función debe devolver una excepción para valores negativos'
    print('OK')

test_previous_square()

```

Output:
```
OK
```