# Rectangulos

Teniendo en cuenta que un rectángulo se puede representar en un plano a  partir de cuatro puntos y un punto se puede representar con 2 coordenadas, considera las siguientes cuestiones:


Escribe una función que nos indique si cuatro puntos pueden ser un rectángulo.

<img src='rectangulo.svg' alt='rectangulo'>

Para ver si 4 puntos forman un rectángulo en plano es necesario ver que que los lados opuestos son iguales y las 2 diagonales miden lo mismo.
Recuerda de la distancia entre los puntos $(a_x, a_y)$ y $(b_x,b_y)$ es $\sqrt{(b_x-a_x)^2 + (b_y-a_y)^2}$

<!-- grade_id: cell-41368ca9833b2e53, points: 0, grade: false, solution: true -->
```python
import math 
def is_rectangle(ax: float, ay: float, bx: float, by: float, 
                 cx: float, cy: float, dx: float, dy: float) -> bool:
    ### BEGIN SOLUTION
    dab = math.sqrt((ax - bx)**2 + (ay - by)**2)
    dac = math.sqrt((ax - cx)**2 + (ay - cy)**2)
    dad = math.sqrt((ax - dx)**2 + (ay - dy)**2)
    dbc = math.sqrt((bx - cx)**2 + (by - cy)**2)
    dbd = math.sqrt((bx - dx)**2 + (by - dy)**2)
    dcd = math.sqrt((cx - dx)**2 + (cy - dy)**2)

    return math.isclose(dab, dcd) and \
        math.isclose(dac, dbd) and \
        math.isclose(dad, dbc)
    ### END SOLUTION

```

```python
### Escribe aquí tus pruebas
```

<!-- grade_id: cell-6afeec483d67fd48, points: 1, grade: true, solution: false -->
```python
def test_is_rectangle():
    import random
    for _ in range(10):
        ax, ay = random.random()*10, random.random()*10
        vx, vy = random.random()*10, random.random()*10
        fac = random.random()*10
        ux, uy = -fac*vy, fac*vx
        bx, by = ax + vx, ay + vy
        cx, cy = ax + ux, ay + uy
        dx, dy = bx + ux, by + uy
        assert is_rectangle(ax, ay, bx, by, cx, cy, dx, dy),\
            f'Error los puntos (́{ax:6.4f}, {ay:6.4f}) (́{bx:6.4f}, {by:6.4f}) (́{cx:6.4f}, {cy:6.4f}) y  (́{dx:6.4f}, {dy:6.4f}) forman un rectángulo'
        ax += 10
        assert not is_rectangle(ax, ay, bx, by, cx, cy, dx, dy),\
            f'Error los puntos (́{ax:6.4f}, {ay:6.4f}) (́{bx:6.4f}, {by:6.4f}) (́{cx:6.4f}, {cy:6.4f}) y  (́{dx:6.4f}, {dy:6.4f}) NO forman un rectángulo'
        ax -= 10

        ay -= 10
        assert not is_rectangle(ax, ay, bx, by, cx, cy, dx, dy),\
            f'Error los puntos (́{ax:6.4f}, {ay:6.4f}) (́{bx:6.4f}, {by:6.4f}) (́{cx:6.4f}, {cy:6.4f}) y  (́{dx:6.4f}, {dy:6.4f}) NO forman un rectángulo'
        ay += 10
    print('OK')

test_is_rectangle()
```

Output:
```
OK
```

Escribe una función que calucle el centro de un rectángulo

<!-- grade_id: cell-ea4872ee5ec3b63f, points: 0, grade: false, solution: true -->
```python
def center(ax: float, ay: float, bx: float, by: float, 
           cx: float, cy: float, dx: float, dy: float) -> float:
    ### BEGIN SOLUTION
    assert is_rectangle(ax,ay, bx,by, cx,cy, dx,dy), f'Error, los puntos (́{ax:6.4f}, {ay:6.4f}) (́{bx:6.4f}, {by:6.4f}) (́{cx:6.4f}, {cy:6.4f}) y  (́{dx:6.4f}, {dy:6.4f}) no forman rectángulo'
    centerx = (ax+cx)/2
    centery = (ay+cy)/2   
    ### END SOLUTION
    return centerx, centery

```

```python
### Escribe aquí tus pruebas
```

<!-- grade_id: cell-3a1e2fc46ba5a7c9, points: 1, grade: true, solution: false -->
```python
def test_center():
    import random
    for _ in range(10):
        centerx, centery = random.random()*10, random.random()*10
        vx, vy = random.random()*10, random.random()*10
        fac = random.random()*10
        ux, uy = -fac*vy, fac*vx

        ax, ay = centerx - (vx + ux)/2, centery - (vy + uy)/2
        bx, by = ax + vx, ay + vy
        cx, cy = bx + ux, by + uy
        dx, dy = ax + ux, ay + uy

        rcx, rcy = center(ax, ay, bx, by, cx, cy, dx, dy)
        assert math.isclose(centerx, rcx) and math.isclose(centery, rcy), \
            f"""
El centro de (́{ax:6.4f}, {ay:6.4f}) (́{bx:6.4f}, {by:6.4f}) (́{cx:6.4f}, {cy:6.4f}) y  (́{dx:6.4f}, {dy:6.4f})
es (́{centerx:6.4f}, {centery:6.4f}) no (́{rcx:6.4f}, {rcy:6.4f})
"""
    print('OK')

test_center()

```

Output:
```
OK
```