# Los siete mensajeros

El relato de Los Siete Mensajeros de Dino Buzzati (1906--1972) 
contiene el siguiente fragmento:

Partí a explorar el reino de mi padre, pero día a día me alejo
más de la ciudad y las noticias que me llegan se hacen cada
vez más escasas. [...]

Aunque despreocupado ---¡mucho más de lo que lo soy ahora!---, pensé
en el modo de comunicarme con mis allegados y, de entre los 
caballeros de mi escolta, elegí a los siete mejores para que me 
hicieran de mensajeros. [...]

Poco habituado a estar lejos de casa mandé al primero, Alejandro,
la noche del segundo día de viaje, cuando habíamos recorrido ya
unas ochenta leguas. Para asegurarme la continuidad de las 
comunicaciones, la noche siguiente envié al segundo, luego al
tercero, luego al cuarto, y así de forma consecutiva hasta la
octava noche del viaje, en que partió Gregorio. El primero aún
no había vuelto.

Éste nos alcanzó la décima noche, mientras nos hallábamos plantando
el campamento para pernoctar en un valle deshabitado. Supe por
Alejandro que su rapidez había sido inferior a la prevista; yo había
pensado que, yendo solo y montando un magnífico corcel, podría recorrer
en el mismo tiempo el doble de distancia que nosotros; sin embargo,
sólo había podido recorrer la equivalente a una vez y media; [...]
lo mismo ocurrió con los demás. Bartolomé, que partió hacia
la ciudad la tercera noche de viaje, volvió a la decimoquinta. Cayo, 
que partió la cuarta, no regresó hasta la vigésima.


Escribe una función que permita calcular cuándo regresará un 
mensajero conociendo el día del viaje en el que se encuentra 
en el momento de su partida. Es decir, conocido el día de partida de 
un mensajero,
la función deberá anunciar el día de su llegada.

Si no te sale la solución puedes leer el relato completo para encontrarla.


<!-- grade_id: cell-f9d40436ad82ab68, points: 0, grade: false, solution: true -->
```python
def mensajeros(day: int) -> int:
    ### BEGIN SOLUTION
    return 5*day
    ### END SOLUTION
```

```python
### Escribe aquí tus pruebas
```

<!-- grade_id: cell-e387d26675191cb9, points: 1, grade: true, solution: false -->
```python
def test_mensajeros():
    tests = [
        (1, 5),
        (2, 10),
        (5, 25)
    ]
    for num, res in tests:
        r = mensajeros(num)
        print(f'{num}, {r}.....', end='')
        assert res == r, f'Error, debería ser {res}'
        print('OK')
        
test_mensajeros()
```

Output:
```
1, 5.....OK
2, 10.....OK
5, 25.....OK
```