# El beso exacto

Como puedes ver en el dibujo, 
dados tres círculos, $C_1$, $C_2$ y $C_3$, tangentes entre sí dos a dos, 
existen exactamente dos círculos que son tangentes a los tres anteriores:

<img src="soddy.svg" alt="Círculos tangentes, y los tangentes comunes">


Hay una fórmula que permite calcular el radio de estos círculos 
tangentes:    
si tenemos cuatro círculos $C_1$, $C_2$, $C_3$ y $C_4$, 
tangentes entre sí, con radios $r_1$, $r_2$, $r_3$ y $r_4$,
respectivamente, entonces se verifica la fórmula

$2(s_1^2 + s_2^2 + s_3^2 + s_4^2) = (s_1 + s_2 + s_3 + s_4)^2$

donde $s_i = 1/r_i$.
No hay magia alguna; de esta fórmula resultan dos soluciones
porque es una ecuación de segundo grado.
La solución de menor valor absoluto siempre es positiva, da el radio del círculo menor. La otra  es negativa si el círculo al que corresponde encierra a los otros tres. Si aparece como solución $s_4=0$ (o sea, $r_4=\infty$), tendremos que el círculo tangente correspondiente se ha *convertido en una recta.*

<em>
Four circles to the kissing come.<br>     
The smaller are the benter.          <br> 
The bend is just the inverse of       <br>
The distance form the center.         <br>
Though their intrigue left Euclid dumb<br>
There's now no need for rule of thumb.<br>
Since zero bend's a dead straight line<br>
And concave bends have minus sign,    <br>
The sum of the squares of all four bends<br>
Is half the square of their sum.<br>
</em>

> Frederick Soddy (1877-1956)

Realiza una función llamada `soddy` que devuelca el radio de las circunferencias menor y mayor tangentes a 2 cirunferencias dadas. Si alguna de los radios no es positivo se debe lanzar una excepción.

<!-- grade_id: cell-090ba3a6bcb4b4e0, points: 0, grade: false, solution: true -->
```python
import math

def soddy(r1:float, r2:float, r3:float) -> float:
    ### BEGIN SOLUTION
    assert r1>0 and r2>0 and r3>0

    s1 = 1.0/r1
    s2 = 1.0/r2
    s3 = 1.0/r3


    a = s1*s1 + s2*s2 + s3*s3
    b = s1 + s2 + s3

    sol1 = b + math.sqrt(2*(b*b - a))
    sol2 = abs(b - math.sqrt(2*(b*b - a)))
    
    small_r = 1.0 / sol1
    big_r = 1.0 / sol2 

    ### END SOLUTION
    return small_r, big_r
```

```python
### Aquí tus pruebas

```

<!-- grade_id: cell-7b3eea2f4b777c09, points: 4, grade: true, solution: false -->
```python
def test_soddy():
    import math
    tests = [
        ((1,2,3), (0.2608695652173913, 6)),
        ((1,1,1), (0.15470053837925155, 2.1547005383792524)),
        ((2,2,2), (0.3094010767585031, 4.309401076758505))
    ]

    for test, result in tests:
        r = soddy(*test)
        print(f'{test}, {r}....', end='')
        assert math.isclose(r[0], result[0]) and math.isclose(r[1], result[1]), f'El resultado debería ser {result}'
        print('OK')

    tests = [
        (0,2,3),
        (1,0,1),
        (2,2,0)
    ]
    for test in tests:
        print(f'{test}....', end='')
        try:
            r = soddy(*test)
        except ZeroDivisionError:
            assert False, f'Error no encontrado {test}'
        except AssertionError:
            pass
        print('OK')


    tests = [
        (-1,2,3),
        (1,-1,1),
        (2,2,-1)
    ]

    for test in tests:
        print(f'{test}....', end='')
        try:
            r = soddy(*test)
        except ValueError:
            assert False, f'Error no encontrado ́{test}'
        except AssertionError:
            pass

        print('OK')
            
    
test_soddy()
```

Output:
```
(1, 2, 3), (0.2608695652173913, 6.000000000000005)....OK
(1, 1, 1), (0.15470053837925155, 2.1547005383792524)....OK
(2, 2, 2), (0.3094010767585031, 4.309401076758505)....OK
(0, 2, 3)....OK
(1, 0, 1)....OK
(2, 2, 0)....OK
(-1, 2, 3)....OK
(1, -1, 1)....OK
(2, 2, -1)....OK
```