# Descomposición de números

En este ejercicio proponemos la definición de funciones sencillas 
que permitan descomponer un número de múltiples formas.      

Escribe una función que devuelva
la parte más significativa desde el $n$-ésimo dígito de un número.

<img style='height:8em' src='digitos1.png'>

La función debe llamarse `most_significant`. El primer argumento debe ser un `int` conteniendo el número y el segundo otro `int` conteniendo la posición.
No se deben usar ni potencias ni la función `str`.

<!-- grade_id: cell-e563669693c0f657, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def most_significant(number:int, position:int, base:int=10) ->int:
    """
    This function returns the most significant part
    of a number in the given base from position
    Example
    number = 351372
    pos = 4
    base = 10
    retruns 35
    

    """
    i = 0
    while i<position:
        number = number//base
        i += 1
    return number
### END SOLUTION
```

```python

```

<!-- grade_id: cell-30629d172f535e36, points: 1, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        print(f'"{test}", "{ex_res}"....', end='')
        res = fun(*test)
        assert res == ex_res, f'Error, the result is "{res}"'
        print('OK.')

        
def test_most_significant():
    wrapper =  most_significant   
    tests = [
        ((351372, 0), 351372),        
        ((351372, 1), 35137),        
        ((351372, 2), 3513),        
        ((351372, 3), 351),        
        ((351372, 4), 35),        
        ((351372, 6), 0),        
        ((351372, 7), 0),        
    ]
    runtests(tests, wrapper)
    
test_most_significant()
```

Output:
```
"(351372, 0)", "351372"....OK.
"(351372, 1)", "35137"....OK.
"(351372, 2)", "3513"....OK.
"(351372, 3)", "351"....OK.
"(351372, 4)", "35"....OK.
"(351372, 6)", "0"....OK.
"(351372, 7)", "0"....OK.
```

Escribe una función que devuelva la parte menos significativa hasta el $n$-ésimo dígito de un número.

<img style='height:8em' src='digitos2.png'>

La función debe llamarse `less_significant`. El primer argumento debe ser un `int` conteniendo el número y el segundo otro `int` conteniendo la posición.
No se deben usar ni potencias ni la función `str`.

<!-- grade_id: cell-a36042eb6e9a89c2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def less_significant(number: int, position:int, base:int=10) -> int:
    """
    This function returns the least significant part
    of a number in the given base from position. 
    
    Example
    number = 351372
    pos = 4
    base = 10
    retruns 372
    """
    i = 0
    part = 0
    pot = 1
    while i<position-1:
        digit = number%base
        part = digit*pot + part
        number = number//base
        pot *= base
        i += 1
    return part
### END SOLUTION

```

```python

```

<!-- grade_id: cell-08d7f5c27d77010c, points: 1, grade: true, solution: false -->
```python
def test_less_significant():
    wrapper =  less_significant
    tests = [
        ((351372, 0), 0),        
        ((351372, 1), 0),        
        ((351372, 2), 2),        
        ((351372, 3), 72),        
        ((351372, 4), 372),        
        ((351372, 5), 1372),        
        ((351372, 6), 51372),        
        ((351372, 8), 351372),        
    ]
    runtests(tests, wrapper)
    
test_less_significant()
```

Output:
```
"(351372, 0)", "0"....OK.
"(351372, 1)", "0"....OK.
"(351372, 2)", "2"....OK.
"(351372, 3)", "72"....OK.
"(351372, 4)", "372"....OK.
"(351372, 5)", "1372"....OK.
"(351372, 6)", "51372"....OK.
"(351372, 8)", "351372"....OK.
```

Escribe una función que nos devuelva el dígito $n$-ésimo de un número.

<img style='height:10em' src='digitos3.png'>

La función debe llamarse `less_significant`. El primer argumento debe ser un `int` conteniendo el número y el segundo otro `int` conteniendo la posición.
No se deben usar ni potencias ni la función `str`. La función debe controla que la posición sea positiva (la primera posición es la 1).

<!-- grade_id: cell-670b4f7a02d78182, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def digit_in_pos(number: int, position:int, base:int=10) -> int:
    """
    This function returns the digit 
    of a number in the given base in the given position. 
    
    Example
    number = 351372
    pos = 4
    base = 10
    retruns 1
    """
    assert position>0, 'The position must be a positive integer'
    i = 0
    while i<position:
        digit = number%base
        number = number//base
        i += 1
    return digit
### END SOLUTION

```

```python

```

<!-- grade_id: cell-859672734b482eee, points: 1, grade: true, solution: false -->
```python
def test_digit_in_pos():
    wrapper =  digit_in_pos
    tests = [
        ((351372, 1), 2),        
        ((351372, 2), 7),        
        ((351372, 3), 3),        
        ((351372, 4), 1),        
        ((351372, 5), 5),        
        ((351372, 6), 3),        
    ]
    runtests(tests, wrapper)
    
test_digit_in_pos()
```

Output:
```
"(351372, 1)", "2"....OK.
"(351372, 2)", "7"....OK.
"(351372, 3)", "3"....OK.
"(351372, 4)", "1"....OK.
"(351372, 5)", "5"....OK.
"(351372, 6)", "3"....OK.
```