# Conjetura Goldbach

En una carta escrita a Leonhard Euler en 1742, 
Christian Goldbach (1690-1764) 
afirmó (sin demostrarlo) que todo número par mayor que 2 es la suma
de dos números primos.
Para poner a prueba esta conjetura
(hasta cierto punto, claro está), basta con avanzar
a través de los primeros $n$ números pares hasta encontrar uno que
no verifica esa propiedad, o hasta llegar al último, ratificándose
la conjetura hasta ese punto.

Realiza una función `decompose` que dado un número par, devuelva dos números primos cuya suma sea el número dado. Puesto que nadie haya demostrado que ambos números existan la función debe devolver un booleano indicando que estos números existen. En difinitiva la función devuelve 3 valores: el primero es el booleano y los dos siguientes son dos enteros, si el primero es `True` entonces los dos son números son primos ordenados y su suma es la suma del valor de entrada.

Necesitarás usar la función `is_prime` que compruebe si un entero mayor que 1 es primo o no. 

<!-- grade_id: cell-9f67b1dec2f403cd, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def is_prime(n:int) -> bool:
    i = 2
    while (i*i <= n) and ((n%i) != 0):
        i = i + 1
    return i*i > n

def decompose(n):
    assert n>2
    a, b = 2, n - 2
    # inv : n = a + b
    # cota : b - a
    while ((a <= b) and not (is_prime(a) and is_prime(b))) :
        # a < b is a prune in search space due to commutative. (on conjecture always true)
        a, b = a + 1, b - 1
    
    value = a <= b
    return value, a, b 

### END SOLUTION
```

```python

```

<!-- grade_id: cell-2594b859c93e1f7e, points: 1, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        print(f'"{test}", "{ex_res}"....', end='')
        res = fun(*test)
        assert res == ex_res, f'Error, the result is "{res}"'
        print('OK.')

        
def test_decompose():
    wrapper = decompose    
    tests = [
        ((4,), (True, 2, 2)),
        ((6,), (True, 3, 3)),
        ((8,), (True, 3, 5)),
        ((10,), (True, 3, 7)),
        ((12,), (True, 5, 7)),
        ((14,), (True, 3, 11)),
        ((16,), (True, 3, 13)),
        ((18,), (True, 5, 13)),
        ((20,), (True, 3, 17)),
        ((22,), (True, 3, 19)),
        ((24,), (True, 5, 19)),
        ((26,), (True, 3, 23)),
        ((28,), (True, 5, 23)),
        ((30,), (True, 7, 23)),
        ((32,), (True, 3, 29)),
        ((34,), (True, 3, 31)),
        ((36,), (True, 5, 31)),
        ((38,), (True, 7, 31)),
        ((40,), (True, 3, 37)),
        ((42,), (True, 5, 37)),
        ((44,), (True, 3, 41)),
        ((46,), (True, 3, 43)),
        ((48,), (True, 5, 43)),
        ((50,), (True, 3, 47)),
        ((52,), (True, 5, 47)),
        ((54,), (True, 7, 47)),
        ((56,), (True, 3, 53)),
        ((58,), (True, 5, 53)),
        ((60,), (True, 7, 53)),
        ((62,), (True, 3, 59)),
        ((64,), (True, 3, 61)),
        ((66,), (True, 5, 61)),
        ((68,), (True, 7, 61)),
        ((70,), (True, 3, 67)),
        ((72,), (True, 5, 67)),
        ((74,), (True, 3, 71)),
        ((76,), (True, 3, 73)),
        ((78,), (True, 5, 73)),
        ((80,), (True, 7, 73)),
        ((82,), (True, 3, 79)),
        ((84,), (True, 5, 79)),
        ((86,), (True, 3, 83)),
        ((88,), (True, 5, 83)),
        ((90,), (True, 7, 83)),
        ((92,), (True, 3, 89)),
        ((94,), (True, 5, 89)),
        ((96,), (True, 7, 89)),
        ((98,), (True, 19, 79)),
    ]
    runtests(tests, wrapper)
    
test_decompose()
```

Output:
```
"(4,)", "(True, 2, 2)"....OK.
"(6,)", "(True, 3, 3)"....OK.
"(8,)", "(True, 3, 5)"....OK.
"(10,)", "(True, 3, 7)"....OK.
"(12,)", "(True, 5, 7)"....OK.
"(14,)", "(True, 3, 11)"....OK.
"(16,)", "(True, 3, 13)"....OK.
"(18,)", "(True, 5, 13)"....OK.
"(20,)", "(True, 3, 17)"....OK.
"(22,)", "(True, 3, 19)"....OK.
"(24,)", "(True, 5, 19)"....OK.
"(26,)", "(True, 3, 23)"....OK.
"(28,)", "(True, 5, 23)"....OK.
"(30,)", "(True, 7, 23)"....OK.
"(32,)", "(True, 3, 29)"....OK.
"(34,)", "(True, 3, 31)"....OK.
"(36,)", "(True, 5, 31)"....OK.
"(38,)", "(True, 7, 31)"....OK.
"(40,)", "(True, 3, 37)"....OK.
"(42,)", "(True, 5, 37)"....OK.
"(44,)", "(True, 3, 41)"....OK.
"(46,)", "(True, 3, 43)"....OK.
"(48,)", "(True, 5, 43)"....OK.
"(50,)", "(True, 3, 47)"....OK.
"(52,)", "(True, 5, 47)"....OK.
"(54,)", "(True, 7, 47)"....OK.
"(56,)", "(True, 3, 53)"....OK.
"(58,)", "(True, 5, 53)"....OK.
"(60,)", "(True, 7, 53)"....OK.
"(62,)", "(True, 3, 59)"....OK.
"(64,)", "(True, 3, 61)"....OK.
"(66,)", "(True,
```