# Palíndromos

Una palabra es un palíndromo se se lee igual del derecho que del revés. Por ejemplo `anilina`. Haz una función llamada `is_palindrome` que indique si la palabra pasada como argumento es un palíndromo.

<!-- grade_id: cell-82e9a011e7f1a11e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def is_palindrome(s: str) -> bool:
    i = 0
    while i<len(s)//2 and s[i]==s[-i-1]:
        i = i + 1
    return i>=len(s)//2
### END SOLUTION
```

```python

```

Output:
```
True
```

<!-- grade_id: cell-19e8c742e23bbc99, points: 1, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        print(f'"{test}", "{ex_res}"....', end='')
        res = fun(*test)
        assert res == ex_res, f'Error, the result is "{res}"'
        print('OK.')

        
def test_is_palindrome():
    def wrapper(s: str):
        return is_palindrome(s)
    
    tests = [
        (('', ), True),
        (('a', ), True),        
        (('abc', ), False),
        (('abcba', ), True),
        (('abcca', ), False),
        (('abba', ), True),
        (('abca', ), False),
    ]
    runtests(tests, wrapper)


    
test_is_palindrome()    
```

Output:
```
"('',)", "True"....OK.
"('a',)", "True"....OK.
"('abc',)", "False"....OK.
"('abcba',)", "True"....OK.
"('abcca',)", "False"....OK.
"('abba',)", "True"....OK.
"('abca',)", "False"....OK.
```

Realiza una función que indique si la representación en una base dada es capicúa. La función se debe llamar `is_palindrome_number`. Debe admitir dos parámetros, uno es el número en sí, y otro es la base. El número debe ser no negativo y la base debe ser un entero entre 2 y 10 (ambos incluidos).

Ejemplo
  252 en base 10 es capicúa.
  252 en base  2 es 11111100 no es capicúa
  5 en base 2: 101 es capicúa.
  
**PISTA:** Debes usar una función que dado un entero, devuelva su representación en la base deseada como cadena de caracteres (vista en clase).

<!-- grade_id: cell-75f0bd443bda3a99, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def repr_int(num: int, base:int):
    rpr = ''
    while num>0:
        digit = num % base
        rpr = chr(digit + ord('0')) + rpr
        num = num // base
    if rpr == '':
        rpr = '0'
    return rpr


def is_palindrome_number(num: int, base:int) -> bool:
    assert base>1 
    assert base<11
    assert num>=0
    
    return is_palindrome(repr_int(num, base))
### END SOLUTION
```

```python

```

<!-- grade_id: cell-b9171c9833368745, points: 1, grade: true, solution: false -->
```python
def test_is_palindrome_number():
    def wrapper(num: int, base:int):
        return is_palindrome_number(num, base)
    
    tests = [
        ((0, 2), True),
        ((2, 10), True),        
        ((5, 2), True),
        ((13, 10), False),
        ((252, 10), True),
        ((252, 2), False)
    ]
    runtests(tests, wrapper)
    
    try:
        s = is_palindrome_number(10,0)
    except AssertionError:
        pass
    else:
        assert False, 'Base 0 not detected'

    try:
        s = is_palindrome_number(10,1)
    except AssertionError:
        pass
    else:
        assert False, 'Base 1 not detected'
        
    try:
        s = is_palindrome_number(10,-1)
    except AssertionError:
        pass
    else:
        assert False, 'Negative base not detected'

    try:
        s = is_palindrome_number(-1,2)
    except AssertionError:
        pass
    else:
        assert False, 'Negative number not detected'


    
test_is_palindrome_number()    
```

Output:
```
"(0, 2)", "True"....OK.
"(2, 10)", "True"....OK.
"(5, 2)", "True"....OK.
"(13, 10)", "False"....OK.
"(252, 10)", "True"....OK.
"(252, 2)", "False"....OK.
```

## Conjetura de la formación de palíndromos.

Dado un número, lo sumamos a su reverso.
Si esta suma es un palíndromo, entonces paramos;
y si no, repetimos el proceso con el número obtenido de dicha suma,
hasta dar con un palíndromo.
            $$
            \begin{array}{ccccccccc}
              59 &\rightarrow& 
              \begin{array}{cr}
                &59\\
                +&95\\ 
                \hline & 154
              \end{array}&\rightarrow&
              \begin{array}{cr}
                &154\\
                +&451\\ 
                \hline & 605
              \end{array}&\rightarrow&
              \begin{array}{cr}
                &605\\
                +&506\\ 
                \hline & 1111
              \end{array}
              &\rightarrow& 1111 \\
            \end{array}
            $$
otro ejemplo
$$
            \begin{array}{ccccccccc}
                317 &\rightarrow& 
                \begin{array}{cc}
                  &317\\
                  +&713\\ 
                  \hline & 1030
                \end{array}&\rightarrow&
                \begin{array}{cc}
                  &1030\\
                  +&0301\\ 
                  \hline & 1331
                \end{array}
                &\rightarrow&1331\\
              \end{array}
                $$

Realiza una función que dado un número, nos indique el palíndromo al que se llega según ese procedimiento. La función se debe llamara `conj_palindrome` y debe tomar como parámetro un entero no negativo y devolver otro entero.
**PISTA**: Debes hacer una función que dado su número, calcule el inverso. Además debes usar la función `is_palindrome_number` definida anteriormente.

<!-- grade_id: cell-9374370921996a4d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def reverse(num:int, base=10) -> int:
    rev = 0
    while num>0:
        digit = num%base
        rev = rev * base + digit 
        num = num//base
    return rev

def conj_palindrome(num: int) -> int:
    assert num>=0
    while not is_palindrome_number(num, 10):
        num = num + reverse(num)
    return num
### END SOLUTION
```

```python

```

<!-- grade_id: cell-e29f294965a4443d, points: 0, grade: true, solution: false -->
```python
def test_conj_palindrome():
    def wrapper(num: int):
        return conj_palindrome(num)
    
    tests = [
        ((0, ), 0),
        ((1, ), 1),
        ((59, ), 1111),        
        ((317, ), 1331),
        ((3452565, ), 89222298),
        ((131, ), 131)
    ]
    runtests(tests, wrapper)
        
    try:
        s = conj_palindrome(-1)
    except AssertionError:
        pass
    else:
        assert False, 'Positive number not checked'
        
test_conj_palindrome()
```

Output:
```
"(0,)", "0"....OK.
"(1,)", "1"....OK.
"(59,)", "1111"....OK.
"(317,)", "1331"....OK.
"(3452565,)", "89222298"....OK.
"(131,)", "131"....OK.
```