Aproximar PI con dardos

<!-- grade_id: cell-d48eca2aa1bcb480, points: 1, grade: true, solution: true -->
```python
import random
def aproximar_pi(dardos):
### BEGIN SOLUTION
    dentro = 0
    i = 0
    while i < dardos:
        x = random.uniform(-1, 1)
        y = random.uniform(-1, 1)
        if x*x + y*y <=1:
            dentro += 1
        i += 1
    return 4*(float(dentro)/dardos)
### END SOLUTION
```

<!-- grade_id: cell-fbb625b87071de7d, points: 1, grade: true, solution: false -->
```python
pi = aproximar_pi(10000)
assert pi>3 and pi<3.2, "Falla, pero al ser un algorimo probabilista, no estoy seguro"
print('Parece que funciona')
```