# Código sustitución monoalfabético

Un código monoalfabético de sustitución es aquél en el que las
letras del alfabeto en que se escribe el mensaje original 
son sustituidas por otras letras del mismo alfabeto. Es decir,
dado un alfabeto, calculamos una permutación suya (que llamaremos criptgrama)
y, para cifrar un mensaje, simplemente
cambiamos las letras según nos indique esa permutación.

|A|B|C|D|E|F|G|H|I|J|K|L|M|N|O|P|Q|R|S|T|U|V|W|X|Y|Z|
|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|
|J|W|K|U|L|Z|B|P|N|H|E|C|G|Y|V|R|M|Q|S|T|D|A|X|I|F|O|



Con este código, para cifrar la palabra "DARDO", tenemos que consultar
en la primera fila cada una de las letras de la palabra y escribir la
letra que se encuentra en la fila de abajo: para la "D"
escribimos "U", para la "A" escribimos "J", y así sucesivamente hasta
obtener "UJSUQ". Observa que las dos apariciones de la misma 
letra "D" en la palabra "DARDO" son codificadas por la misma letra "U".

Supondremos que los textos a codificar/decodificar están formados por letras mayúsculas entre la "A" y ls "Z", no hay ni tildes ni "Ñ" y las palabras están separadas por blancos (que no se codifican). Según el criptograma anterior, la frase 'HOY ES LUNES' se codifica como 'PVFLSCDNLS'.



**Pista:** observa que para representar una permutacin cualquiera en python basta con considerar una cadena de caracteres de longigud 26. De esa forma
la últiva permutación se puede representar por la cadena de caracteres `JWKULZBPNHECGYVRMQSTDAXIFO`

```python
len('CRIPTOGAMBDEFHJKLNQSUVWXYZ')
```

Output:
```
26
```

## Cifrado
Haz una función que cifre un texto dado. La función se debe llamar `crypt`. Debe tomar 2 parámetros, el primero es el texto a codificar y el segundo es el criptograma.


<!-- grade_id: cell-dd0f1119ff16b62f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def crypt(txt: str, criptogram:str) -> str:
    """
    Encode the line s according to the given criptogram
    """    
    enc=""
    for c in txt:
        if c!=" ":
            enc+=criptogram[ord(c)-ord("A")]
    return enc
### END SOLUTION
```

```python
cpgm = 'CRIPTOGAMBDEFHJKLNQSUVWXYZ'
txt = 'ESTA ASIGNATURA NO LA APRUEBA NI DIOS'
crypt(txt, cpgm)
```

Output:
```
'TQSCCQMGHCSUNCHJECCKNUTRCHMPMJQ'
```

<!-- grade_id: cell-13ac7d73878612ca, points: 1, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        print(f'"{test}", "{ex_res}"....', end='')
        res = fun(*test)
        assert res == ex_res, f'Error, the result is "{res}"'
        print('OK.')
        
def test_crypt():
    tests = [
        (('HOY ES LUNES', 'JWKULZBPNHECGYVRMQSTDAXIFO'),  'PVFLSCDYLS'),
        (('ESTA ASIGNATURA NO LA APRUEBA NI DIOS', 'CRIPTOGAMBDEFHJKLNQSUVWXYZ'), 'TQSCCQMGHCSUNCHJECCKNUTRCHMPMJQ')
    ]
    runtests(tests, crypt)
    
test_crypt()
```

Output:
```
"('HOY ES LUNES', 'JWKULZBPNHECGYVRMQSTDAXIFO')", "PVFLSCDYLS"....OK.
"('ESTA ASIGNATURA NO LA APRUEBA NI DIOS', 'CRIPTOGAMBDEFHJKLNQSUVWXYZ')", "TQSCCQMGHCSUNCHJECCKNUTRCHMPMJQ"....OK.
```

## Decifrado
Ahora realiza la función inversa, se debe llamar 'decrypt', debe tomar un texto (cifrado) y un criptograma como se indica y debe devolver el texto descifrado.


<!-- grade_id: cell-97f6920eedbe49cf, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def inv_cripto(cpgm: str) -> str:
    c = 'A'
    inv = ''
    while c <= 'Z':
        i = 0
        while i<len(cpgm) and cpgm[i]!=c:
            i = i+1
        inv = inv + chr(i+ord('A'))
        c = chr(ord(c)+1)
    return inv
                

def decrypt(txt, cpgm):
    return crypt(txt, inv_cripto(cpgm))
### END SOLUTION                       
```

<!-- grade_id: cell-c719747580c83fb3, points: 1, grade: true, solution: false -->
```python

```

```python
def test_decrypt():
    tests = [
        (('PVFLSCDYLS', 'JWKULZBPNHECGYVRMQSTDAXIFO'),  'HOYESLUNES'),
        (('TQSCCQMGHCSUNCHJECCKNUTRCHMPMJQ', 'CRIPTOGAMBDEFHJKLNQSUVWXYZ'), 'ESTAASIGNATURANOLAAPRUEBANIDIOS')
    ]
    runtests(tests, decrypt)
    
test_decrypt()
```

Output:
```
"('PVFLSCDYLS', 'JWKULZBPNHECGYVRMQSTDAXIFO')", "HOYESLUNES"....OK.
"('TQSCCQMGHCSUNCHJECCKNUTRCHMPMJQ', 'CRIPTOGAMBDEFHJKLNQSUVWXYZ')", "ESTAASIGNATURANOLAAPRUEBANIDIOS"....OK.
```

## La clave
La clave necesaria en este tipo de cifrados es una permutación del 
alfabeto que vayamos a utilizar. Esta clave es difícil de retener de 
memoria y por tanto se necesita su almacenamiento y gestión. Para 
reducir estos inconvenientes, se puede recurrir a códigos 
pseudoaleatorios. Por ejemplo, a partir de una palabra que se utiliza 
como clave. Supongamos que queremos realizar un código monoalfabético 
cuya clave podamos recordar, elijamos como clave la palabra 
"OCULTA". Una forma posible de definir el código es:

A|B|C|D|E|F|G|H|I|J|K|L|M|N|O|P|Q|R|S|T|U|V|W|X|Y|Z
|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|
O|C|U|L|T|A|B|D|E|F|G|H|I|J|K|M|N|P|Q|R|S|V|W|X|Y|Z

Hemos utilizado la palabra clave para definir el
comienzo de la permutación y luego hemos seguido el orden alfabético.
Si la palabra clave tiene letras repetidas, entonces basta con no 
escribirlas de nuevo. Por ejemplo si consideramos "CRIPTOGRAMA" como 
palabra clave, entonces el cifrado quedaría
A|B|C|D|E|F|G|H|I|J|K|L|M|N|O|P|Q|R|S|T|U|V|W|X|Y|Z
|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|
C|R|I|P|T|O|G|A|M|B|D|E|F|H|J|K|L|N|Q|S|U|V|W|X|Y|Z



Escribe una función llamada `no_rep` que tenga como parámetro una "palabra" y que
construya otra "palbr" (aunque no exista en el diccionario) que
sea idéntica a la primera pero que elimine las repeticiones de letras.


<!-- grade_id: cell-6e8e19c7fe048cb5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def no_rep(word: str) -> str:
    new = ''
    for c in word:
        if not c in new:
            new = new+c
    return new
### END SOLUTION
```

```python

```

<!-- grade_id: cell-975c34ebb62ee63f, points: 1, grade: true, solution: false -->
```python
def test_no_rep():
    tests = [
        (('CRIPTOGRAMA',), 'CRIPTOGAM'),
        (('PALABRA',), 'PALBR')
    ]
    runtests(tests, no_rep)
    
test_no_rep()
```

Output:
```
"('CRIPTOGRAMA',)", "CRIPTOGAM"....OK.
"('PALABRA',)", "PALBR"....OK.
```

Escribe una función llamada `gen_crytopgram` que construya una permutación de  un
alfabeto basada en una palabra clave. Es decir, una permutación que
comience con las letras no repetidas de la palabra y continue con las
letras restantes en el orden alfabético.

<!-- grade_id: cell-68f16a4442ac3886, points: 0, grade: false, solution: true -->
```python
###BEGIN SOLUTION
def gen_crytopgram(code: str) -> str:
    cpgm = no_rep(code)
    c = 'A'
    while c<='Z':
        if c not in cpgm:
            cpgm = cpgm + c
        c = chr(ord(c)+1)
    return cpgm
###END SOLUTION
```

```python

```

<!-- grade_id: cell-86b07e8c71c4dfed, points: 1, grade: true, solution: false -->
```python
def test_gen_crytopgram():
    tests = [
        (('OCULTA',),'OCULTABDEFGHIJKMNPQRSVWXYZ'),
        (('CRIPTOGRAMA',),'CRIPTOGAMBDEFHJKLNQSUVWXYZ')
    ]
    runtests(tests, gen_crytopgram)
    
test_gen_crytopgram()
```

Output:
```
"('OCULTA',)", "OCULTABDEFGHIJKMNPQRSVWXYZ"....OK.
"('CRIPTOGRAMA',)", "CRIPTOGAMBDEFHJKLNQSUVWXYZ"....OK.
```

### Cifrado con clave

Haz una función de cifrado `cyrpt_key` que use una clave. A partir de la clave se genera el criptograma y se cofifica el texto con él. El primer parámetro de la función es el texto a cifrar y el segundo la clave. 


<!-- grade_id: cell-1ebee0405cfde98d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def crypt_key(text:str, key:str) -> str:
    cpgm = gen_crytopgram(key)
    return crypt(text, cpgm)
### END SOLUTION
```

```python

```

<!-- grade_id: cell-373dc4c7371c5471, points: 1, grade: true, solution: false -->
```python
def test_crypt_key():
    tests = [
        (('EN UN LUGAR DE LA MANCHA DE CUYO NOMBRE NO QUIERO ACORDARME','PERNAMBUCO'), 'AHTHFTBPLNAFPGPHRUPNARTYIHIGELAHIKTCALIPRILNPLGA'),
        (('EN UN LUGAR DE LA MANCHA DE CUYO NOMBRE NO QUIERO ACORDARME','AVION'), 'NKTKHTCAQONHAJAKIDAONITYLKLJVQNKLPTENQLAILQOAQJN'),
        (('EN UN LUGAR DE LA MANCHA DE CUYO NOMBRE NO QUIERO ACORDARME','PI'), 'CMUMKUEPRBCKPLPMAFPBCAUYNMNLIRCMNQUGCRNPANRBPRLC'),
    ]
    runtests(tests, crypt_key)
test_crypt_key()
```

Output:
```
"('EN UN LUGAR DE LA MANCHA DE CUYO NOMBRE NO QUIERO ACORDARME', 'PERNAMBUCO')", "AHTHFTBPLNAFPGPHRUPNARTYIHIGELAHIKTCALIPRILNPLGA"....OK.
"('EN UN LUGAR DE LA MANCHA DE CUYO NOMBRE NO QUIERO ACORDARME', 'AVION')", "NKTKHTCAQONHAJAKIDAONITYLKLJVQNKLPTENQLAILQOAQJN"....OK.
"('EN UN LUGAR DE LA MANCHA DE CUYO NOMBRE NO QUIERO ACORDARME', 'PI')", "CMUMKUEPRBCKPLPMAFPBCAUYNMNLIRCMNQUGCRNPANRBPRLC"....OK.
```

Escribe una función `decrypt_key` que decifre un texto dada una clave. El primer argumento es el texto cifrado y el segundo es la clave.

<!-- grade_id: cell-f57ed53bf1275909, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def decrypt_key(text:str, key:str) -> str:
    cpgm = gen_crytopgram(key)
    return decrypt(text, cpgm)
### END SOLUTION
```

```python

```

<!-- grade_id: cell-d395b2f8ddb8ad1e, points: 1, grade: true, solution: false -->
```python
def test_decrypt_key():
    tests = [
        (('AHTHFTBPLNAFPGPHRUPNARTYIHIGELAHIKTCALIPRILNPLGA','PERNAMBUCO'), 'ENUNLUGARDELAMANCHADECUYONOMBRENOQUIEROACORDARME'),
        (('NKTKHTCAQONHAJAKIDAONITYLKLJVQNKLPTENQLAILQOAQJN','AVION'), 'ENUNLUGARDELAMANCHADECUYONOMBRENOQUIEROACORDARME'),
        (('CMUMKUEPRBCKPLPMAFPBCAUYNMNLIRCMNQUGCRNPANRBPRLC','PI'), 'ENUNLUGARDELAMANCHADECUYONOMBRENOQUIEROACORDARME'),
    ]
    runtests(tests, decrypt_key)
test_decrypt_key()
```

Output:
```
"('AHTHFTBPLNAFPGPHRUPNARTYIHIGELAHIKTCALIPRILNPLGA', 'PERNAMBUCO')", "ENUNLUGARDELAMANCHADECUYONOMBRENOQUIEROACORDARME"....OK.
"('NKTKHTCAQONHAJAKIDAONITYLKLJVQNKLPTENQLAILQOAQJN', 'AVION')", "ENUNLUGARDELAMANCHADECUYONOMBRENOQUIEROACORDARME"....OK.
"('CMUMKUEPRBCKPLPMAFPBCAUYNMNLIRCMNQUGCRNPANRBPRLC', 'PI')", "ENUNLUGARDELAMANCHADECUYONOMBRENOQUIEROACORDARME"....OK.
```