Escribe un función que indique la actividad mas apropiada para una temperatura teniendo en cuenta los siguientes criterios:


* Natación si $t>30$ (`swimming`)
* Tenis si  $20< t<\leq  30$  (`tennis`)
* Golf si $10<t\leq 20$ (`golf`)
* Esquí si $5<t\leq 10$  (`ski`)
* Parchís si $t\leq 5$ (`parchis`)

La función debe devolver la cadena de caracteres indicada entre paréntesis.

<!-- grade_id: cell-a80f4451c8f37f60, points: 0, grade: false, solution: true -->
```python
def activity(temp: float) -> str:
    ### BEGIN SOLUTION
    if temp > 30:
        act = "swimming"
    elif 20 < temp <= 30:
        act = "tennis"
    elif 10 < temp <= 20:
        act = "golf"
    elif 5 < temp <= 10:
        act = "ski"
    else:
        act = "parchis"
    return act
    ### END SOLUTION
```

```python
### escribe aquí tu solución

```

<!-- grade_id: cell-3dad62f0121bd594, points: 0, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        res = fun(*test)
        print(f'{test}, {res}....', end='')
        assert res == ex_res, f'Error, the expexted result is {ex_res}'
        print('OK.')
        
def test_activity():
    tests = [
        ((31,),"swimming"),
        ((30,),"tennis"),
        ((25,),"tennis"),
        ((20,),"golf"),
        ((15,),"golf"),
        ((10,),"ski"),
        ((7,),"ski"),
        ((5,),"parchis"),
        ((2,),"parchis"),
    ]
    runtests(tests, activity)

test_activity()
```

Output:
```
(31,), swimming....OK.
(30,), tennis....OK.
(25,), tennis....OK.
(20,), golf....OK.
(15,), golf....OK.
(10,), ski....OK.
(7,), ski....OK.
(5,), parchis....OK.
(2,), parchis....OK.
```