# Cuadrantes

Escribe una función que nos indique el cuadrante en el que se encuentra un punto en el plano.

<img src='cuadrantes.svg' alt=''>

Si alguna de las componentes es 0 (el punto está en un eje), debe devolver el cuadrante menor que delimita el eje.

<!-- grade_id: cell-d396b81a2b451f50, points: 0, grade: false, solution: true -->
```python
def quadrant(x: float, y:float) -> int:
    #### BEGIN SOLUTION
    if x>=0 and y>=0:
        q = 1
    elif y>=0: # x<0
        q = 2
    elif x<=0: #y<0
        q = 3
    else: #x<0 and y<0
        q = 4
    return q
    #### END SOLUTION
```

```python
### Pruebas
```

<!-- grade_id: cell-c9820ac49790dcf6, points: 1, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        res = fun(*test)
        print(f'{test}, {res}....', end='')
        assert res == ex_res, f'Error, the expexted result is {ex_res}'
        print('OK.')

        
def test_quadrant():
    tests = [
        ((0,0), 1),
        ((1,0), 1),
        ((1, 1), 1),
        ((0, 1), 1),
        ((-1, 1), 2),
        ((-1, 0), 2),
        ((-1, -1), 3),
        ((0, -1), 3),
        ((1, -1), 4)
    ]
    runtests(tests, quadrant)
    
test_quadrant()
```

Output:
```
(0, 0), 1....OK.
(1, 0), 1....OK.
(1, 1), 1....OK.
(0, 1), 1....OK.
(-1, 1), 2....OK.
(-1, 0), 2....OK.
(-1, -1), 3....OK.
(0, -1), 3....OK.
(1, -1), 4....OK.
```