#  ¿Me puede decir la hora?  

Escribe un programa que reciba tres valores de entrada que corresponden a horas, minutos y  
segundos. El programa debe indicar si se trata de una hora correcta o no.

<!-- grade_id: cell-683d3114d4f866c7, points: 0, grade: false, solution: true -->
```python
def is_correct_time(hour:int, minute:int, second:int) -> bool:
    ### BEGIN SOLUTION
    """
    This function checks if the passed time is correct
    @type hour: int
    @type minute: int
    @type secomd: int
    @rtype: bool
    """
    return ( 0<=hour and hour<24 ) and \
        (0<=minute and minute<60 ) and \
        (0<=second and second<60 )
    ### END SOLUTION
    
    
```

```python
### Pruebas
```

<!-- grade_id: cell-4d9a374ac8f23b04, points: 0, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        res = fun(*test)
        print(f'{test}, {res}....', end='')
        assert res == ex_res, f'Error, the expexted result is {ex_res}'
        print('OK.')


def test_is_correct_time():
    tests = [
        ((10, 10, 10), True),
        ((0, 0, 0), True),
        ((23, 59, 59), True),
        ((24, 59, 59), False),
        ((24, 60, 59), False),
        ((24, 59, 60), False),
        ((-1, 59, 59), False),
        ((24, -1, 59), False),
        ((24, 59, -1), False),
    ]
    runtests(tests, is_correct_time)

test_is_correct_time()

```

Output:
```
(10, 10, 10), True....OK.
(0, 0, 0), True....OK.
(23, 59, 59), True....OK.
(24, 59, 59), False....OK.
(24, 60, 59), False....OK.
(24, 59, 60), False....OK.
(-1, 59, 59), False....OK.
(24, -1, 59), False....OK.
(24, 59, -1), False....OK.
```