# Ecuación de segundo grado.

Escribe un programa que calcule la solución de una ecuación de segundo grado $Ax^2 + Bx + C = 0$. 
La función debe devolver las dos soluciones.

La solución debe tener en cuenta que $A$ puede ser $0$, en este caso la función debe debe devolver 2 soluciones la segunda de ellas es `None`.

Si el discriminante es negativo (no hay soluciones reales), los dos valores que se devuelven son `None`.

<!-- grade_id: cell-c0a75aec99d347b2, points: 0, grade: false, solution: true -->
```python
import math

def grade_2_equation(a: float, b: float, c: float) -> (float, float):
    assert not math.isclose(a, 0) or not math.isclose(b, 0), f'Error: there is no equation'
    ### BEGIN SOLUTION
    disc = b*b - 4*a*c
    if disc < 0:
        sol1, sol2 = None, None
    elif math.isclose(a, 0):
        sol1 = -c/b
        sol2 = None
    else:
        sol1 = (-b + math.sqrt(disc)) / (2*a)
        sol2 = (-b - math.sqrt(disc)) / (2*a)
    
    return sol1, sol2
    ### END SOLUTION
```

```python
### Escribe tus pruebas
grade_2_equation(3, 0,18)
```

Output:
```
(None, None)
```

<!-- grade_id: cell-2d8e869ca8cfe038, points: 0, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        res = fun(*test)
        sol1, sol2 = res
        ex_s1, ex_s2 = ex_res
        print(f'{test}, {res}....', end='')
        err_msg = f'Error, the expexted result is {ex_res}'
        if ex_s1 is None and ex_s2 is None:
            assert sol1 is None and sol2 is None, err_msg        
        elif ex_s2 is None:
            assert math.isclose(sol1, ex_s1) and sol2 is None, err_msg
        else:
            assert math.isclose(sol1, ex_s1) and math.isclose(sol2, ex_s2), err_msg            
        print('OK.')

def test_grade_2_equation():
    tests = [
        ((3,-15,18), (3.0, 2.0)),
        ((0, 4,18), (-4.5, None)),
        ((3, 0, 18),  (None, None)),
        ((3, 0, -12), (2,-2)),
        ((2,-12,18), (3,3)),
        ((45, 1, 2), (None, None))
    ]
    runtests(tests, grade_2_equation)
    
test_grade_2_equation()
    
        
```

Output:
```
(3, -15, 18), (3.0, 2.0)....OK.
(0, 4, 18), (-4.5, None)....OK.
(3, 0, 18), (None, None)....OK.
(3, 0, -12), (2.0, -2.0)....OK.
(2, -12, 18), (3.0, 3.0)....OK.
(45, 1, 2), (None, None)....OK.
```