# Rotaciones

Imaginemos las letras del alfabeto ordenadas y dispuestas en círculo.
Esto es, a la derecha de la `A` se encuentra la `B`, luego la `C` y así
sucesivamente hasta la `Z`; a la derecha de la `Z` se encuentra 
nuevamente la `A`.

Definimos una rotación de longitud $n$ como aquella que 
lleva a una determinada letra $n$ posiciones hacia su derecha.

Vamos a considerar letras en mayúsculas sin tildes ni `Ñ`.

## Ejemplo
La rotación de longitud 1 lleva la $A$ a la $B$, la 
$V$ a la $W$ y la $Z$ a la $A$.

La rotación de longitud 3 lleva la $A$ a la $D$, la $V$ a la $Y$ y la
$Z$ a la $C$.



## Rotación de longitud 1

Escribe en python una función que permita calcular la rotación de longitud 1. Es decir,
dada una letra del alfabeto 
la función debería devolver la letra corrrespondiente a su rotación.


<!-- grade_id: cell-cb766c0e262b9d26, points: 0, grade: false, solution: true -->
```python
def rotation1(ltr: str) -> str:
    ### BEGIN SOLUTION
    def _rotation(ltr: str, lrot) -> str:
        nchars = ord('Z') - ord('A') + 1
        num = (ord(ltr) - ord('A') + lrot) % nchars
        return chr(num + ord('A'))
    return _rotation(ltr, 1)
    ### END SOLUTION

```

```python
### Pruebas
```

<!-- grade_id: cell-dbbf06cd9763e2d9, points: 1, grade: true, solution: false -->
```python
def runtests(tests, fun):
    for test, ex_res in tests:
        res = fun(*test)
        print(f'{test}, {res}....', end='')
        assert res == ex_res, f'Error, the expexted result is {ex_res}'
        print('OK.')

def test_rotation1():
    tests = [
        (('A',), 'B'),
        (('L',), 'M'),
        (('Y',), 'Z'),
        (('Z',), 'A'),
    ]
    runtests(tests, rotation1)

test_rotation1()

```

Output:
```
('A',), B....OK.
('L',), M....OK.
('Y',), Z....OK.
('Z',), A....OK.
```

## Rotación de longitud arbitraria

Escribe en python una función que permita calcular la rotación de longitud $n$

<!-- grade_id: cell-3690770508a6802a, points: 0, grade: false, solution: true -->
```python
def rotation(ltr: str, lrot) -> str:
    nchars = ord('Z') - ord('A') + 1
    num = (ord(ltr) - ord('A') + lrot) % nchars
    return chr(num + ord('A'))

```

```python
### Pruebas
```

<!-- grade_id: cell-680781279db82f5e, points: 1, grade: true, solution: false -->
```python
def test_rotation():
    tests = [
        (('A',1), 'B'),
        (('A',2), 'C'),
        (('A',3), 'D'),
        (('L',1), 'M'),
        (('L',2), 'N'),
        (('L',3), 'O'),
        (('Y',1), 'Z'),
        (('Y',2), 'A'),
        (('Y',3), 'B'),
        (('Z',1), 'A'),
        (('Z',2), 'B'),
        (('Z',3), 'C'),
    ]
    runtests(tests, rotation)

test_rotation()

```

Output:
```
('A', 1), B....OK.
('A', 2), C....OK.
('A', 3), D....OK.
('L', 1), M....OK.
('L', 2), N....OK.
('L', 3), O....OK.
('Y', 1), Z....OK.
('Y', 2), A....OK.
('Y', 3), B....OK.
('Z', 1), A....OK.
('Z', 2), B....OK.
('Z', 3), C....OK.
```