# Programación II. Segunda práctica obligatoria
Los __polimorfismos de nucleótido único__ (SNPs) son el tipo de variación genética más habitual entre las personas. Un polimorfismo consiste en el cambio de un nucleótido (adenina, guanina, citosina, timina) en una posición concreta del ADN con respecto a un genoma de referencia, y que se encuentra al menos en el 1 por ciento de la población. La mayoría de estos polimorfismos no tienen efectos en el desarrollo y la salud, pero unos pocos sí pueden tener consecuencias en el fenotipo. El fenotipo es la expresión del genotipo que puede observarse (aunque está parcialmente determinado por el entorno). Estos polimorfismos afectarían a nuestra susceptibilidad a padecer determinadas enfermedades o a metabolizar ciertas sustancias de forma diferente. El objetivo de la práctica es la implementación de una aplicación que facilite el tratamiento de ciertos datos relacionados con algunos polimorfismos de nucleótido único.

```python
# Funciones auxiliares empleadas en las pruebas
def equalitySNP(dict_snp1, dict_snp2):
    equal=set(dict_snp1.keys())==set(dict_snp2.keys())
    if equal: 
        for codeN in dict_snp1.keys():
            if (set([disease.upper() for disease in dict_snp1[codeN]])!=set([disease.upper() for disease in dict_snp2[codeN]])):
                return False
    return equal
def equalitySNPs(dictSNP1, dictSNP2, list_codeRSr):
    keys_dictSNP1=set(dictSNP1.keys())
    keys_dictSNP2=set(dictSNP2.keys())
    for code in list_codeRSr:
        if (code in dictSNP1):
            keys_dictSNP1.discard(code)
        if (code in dictSNP2):
            keys_dictSNP2.discard(code)
    equal=keys_dictSNP1==keys_dictSNP2
    if equal: 
        for codeRS in dictSNP1.keys():
            if (codeRS not in list_codeRSr):
                if not equalitySNP(dictSNP1[codeRS], dictSNP2[codeRS]):
                    return False
    return equal   
```

### Ejercicio 1
Diseña una función llamada `readSNPfile(name)` que, dado el nombre de un fichero con la estructura adecuada, lea su información y devuelva un diccionario que la contenga con el formato especificado en el enunciado.

<!-- grade_id: readSNPfile_function, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: readSNPfile_test, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import copy
print('INICIO de las pruebas (readSNPfile())')
tests=[('rs7495174',{}),('rs713598',{}),('rs2241880',{'C/C':['ENFERMEDAD DE CROHN']}),
       ('rs1051730',{'T/T':['ALCOHOLISMO','CARCINOMA BRONCOPULMONAR']}),
       ('rs16944',{'G/G':['DEPRESION','ESQUIZOFRENIA','TRASTORNO BIPOLAR'],'A/A':['OSTEOARTRITIS']}),
       ('rs6897932',{'C/C': ['ESCLEROSIS MULTIPLE']}),
       ('rs6313',{'C/T':['ARTRITIS REUMATOIDE'],'C/C':['ARTRITIS REUMATOIDE'],'T/T':['DEPRESION','TRASTORNO DE PANICO']})]
my_dictSNP=readSNPfile('dataSNP.txt') 
for (codeRS, solution) in tests:
    my_dictSNP_copy=copy.deepcopy(my_dictSNP)
    print('  codigoRS: {}\t resultado: {}'.format(codeRS, solution))
    resultado=my_dictSNP.get(codeRS, {})
    try:     
        assert equalitySNP(resultado, solution) and equalitySNPs(my_dictSNP, my_dictSNP_copy, [])
    except:
        print('    RESULTADO INCORRECTO codigoRS: {} resultado: {}'.format(codeRS, resultado))
print('FIN de las pruebas (readSNPfile())')
### END TESTS
```

### Ejercicio 2
Implementa una función llamada `addSNPinfo(dictSNP,tupleSNP)` que, dados un diccionario `dictSNP` de polimorfismos y la tupla `tupleSNP` de un SNP, incluya la información del SNP en el diccionario. La función no solo debe incorporar nuevos polimorfismos sino también actualizar los SNPs existentes.

<!-- grade_id: addSNPinfo_function, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: addSNPinfo_test, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import copy
print('INICIO de las pruebas (addSNPinfo())')
tests=[(('rs7495174','C/C',['ENFERMEDAD DE CROHN']),{'C/C':['ENFERMEDAD DE CROHN']}),
       (('rs7495174','C/C',['TABAQUISMO']),{'C/C':['ENFERMEDAD DE CROHN','TABAQUISMO']}),
       (('rs7495174','T/C',['CANCER TESTICULAR','MELANOMA']),{'C/C':['ENFERMEDAD DE CROHN','TABAQUISMO'],'T/C':['CANCER TESTICULAR','MELANOMA']}),
       (('rs7495174','C/C',['TABAQUISMO']),{'C/C':['ENFERMEDAD DE CROHN','TABAQUISMO'],'T/C':['CANCER TESTICULAR','MELANOMA']}),
       (('rs713598','C/C',['ADICCION A LA COCAINA']),{'C/C':['ADICCION A LA COCAINA']}),
       (('rs713598','C/C',['ADICCION A LA COCAINA','ESQUIZOFRENIA']),{'C/C':['ADICCION A LA COCAINA','ESQUIZOFRENIA']}),
       (('rs713598','C/C',[]),{'C/C':['ADICCION A LA COCAINA','ESQUIZOFRENIA']})]
my_dictSNP=readSNPfile('dataSNP.txt') 
list_remove=[]
for (tupleSNP, solution) in tests:
    my_dictSNP_copy=copy.deepcopy(my_dictSNP)
    print('  codigoRS: {}\t resultado: {}'.format(tupleSNP[0], solution))
    list_remove.append(tupleSNP[0])
    addSNPinfo(my_dictSNP, tupleSNP)
    resultado=my_dictSNP.get(tupleSNP[0], {})
    try:     
        assert equalitySNP(resultado, solution) and equalitySNPs(my_dictSNP, my_dictSNP_copy, list_remove)
    except:
        print('    RESULTADO INCORRECTO codigoRS: {} resultado: {}'.format(tupleSNP[0], resultado))
print('FIN de las pruebas (addSNPinfo())')
### END TESTS
```

### Ejercicio 3
Diseña una función llamada `removeSNPinfo(dictSNP,tupleSNP)` que, dados un diccionario `dictSNP` de polimorfismos y la tupla `tupleSNP` de un SNP, elimine los datos del SNP del diccionario.

<!-- grade_id: removeSNPinfo_function, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: removeSNPinfo_test, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import copy
print('INICIO de las pruebas (removeSNPinfo())')
tests=[(('rs16944','G/G',['DEPRESION','ESQUIZOFRENIA']),{'A/A':['OSTEOARTRITIS'],'G/G':['TRASTORNO BIPOLAR']}),
       (('rs16944','G/G',['DEPRESION','TABAQUISMO']),{'A/A':['OSTEOARTRITIS'],'G/G':['TRASTORNO BIPOLAR']}),
       (('rs16944','A/A',['OSTEOARTRITIS']),{'G/G': ['TRASTORNO BIPOLAR']}),
       (('rs16944','G/G',['TRASTORNO BIPOLAR']),{}),
       (('rs4880','A/T',['MESOTELIOMA']),{'T/T':['HEMOCROMATOSIS'],'C/C':['MESOTELIOMA']}),
       (('rs1234','T/T',['TABAQUISMO']),{}),(('rs1051730','T/T',[]),{})]
my_dictSNP=readSNPfile('dataSNP.txt') 
list_remove=[]
for (tupleSNP, solution) in tests:
    my_dictSNP_copy=copy.deepcopy(my_dictSNP)
    print('  codigoRS: {}\t resultado: {}'.format(tupleSNP[0], solution))
    list_remove.append(tupleSNP[0])
    removeSNPinfo(my_dictSNP, tupleSNP)
    resultado=my_dictSNP.get(tupleSNP[0], {})
    try:     
        assert equalitySNP(resultado, solution) and equalitySNPs(my_dictSNP, my_dictSNP_copy, list_remove)
    except:
        print('    RESULTADO INCORRECTO codigoRS: {} resultado: {}'.format(tupleSNP[0], resultado))
print('FIN de las pruebas (removeSNPinfo())')
### END TESTS
```

### Ejercicio 4
Implementa una función llamada `getSNPinfo(dictSNP,lstPat)` que, dados un diccionario `dictSNP` de polimorfismos y una lista `lstPat` de patologías, devuelva una lista de pares de la forma `(codeRS,codeN)` de modo que todas las patologías de `lstPat` estén relacionadas con todos los SNP de la lista resultado.

<!-- grade_id: getSNPinfo_function, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b60e2856bb26a4ed, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import copy
print('INICIO de las pruebas (getSNPinfo())')
tests=[(['ALCOHOLISMO'],[('rs1076560','A/C'),('rs1051730','T/T'),('rs1800497','C/T')]),
       (['FIBROSIS'],[('rs738409','C/C'),('rs2736100','T/T')]),
       (['FIBROSIS PULMONAR'],[('rs35705950','G/T'),('rs35705950','T/T')]),
       (['MELANOMA'],[('rs1805007','T/T'),('rs13181','G/G')]),
       (['MELANOMA','CANCER OVARICO'],[('rs13181','G/G')]),
       (['FIBROSIS','CARCINOMA BRONCOPULMONAR','CANCER TESTICULAR'],[('rs2736100','T/T')]),
       (['ALCOHOLISMO','CARCINOMA'],[])
      ]
my_dictSNP=readSNPfile('dataSNP.txt') 
for (lstPat, solution) in tests:
    my_dictSNP_copy=copy.deepcopy(my_dictSNP)
    print('  patologias: {} resultado: {}'.format(lstPat, solution))
    resultado=getSNPinfo(my_dictSNP, lstPat)
    try:     
        assert set(resultado)==set(solution) and equalitySNPs(my_dictSNP, my_dictSNP_copy, [])
    except:
        print('    RESULTADO INCORRECTO codigoRS: {} resultado: {}'.format(lstPat, resultado))
print('FIN de las pruebas (getSNPinfo())')
### END TESTS
```

### Ejercicio 5
Implementa una función llamada `writeSNPfile(dictSNP,name)` que, dados un diccionario `dictSNP` de polimorfismos y el nombre de un fichero nuevo, almacene en este último el contenido de `dictSNP` respetando el formato especificado en el primer apartado.

<!-- grade_id: writeSNPfile_function, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: writeSNPfile_test, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
print('INICIO de las pruebas (writeSNPfile())')
tests=[('rs7495174',{}),('rs713598',{}),('rs2241880',{'C/C':['ENFERMEDAD DE CROHN']}),
       ('rs1051730',{'T/T':['ALCOHOLISMO','CARCINOMA BRONCOPULMONAR']}),
       ('rs16944',{'G/G':['DEPRESION','ESQUIZOFRENIA','TRASTORNO BIPOLAR'],'A/A':['OSTEOARTRITIS']}),
       ('rs6897932',{'C/C': ['ESCLEROSIS MULTIPLE']}),
       ('rs6313',{'C/T':['ARTRITIS REUMATOIDE'],'C/C':['ARTRITIS REUMATOIDE'],'T/T':['DEPRESION','TRASTORNO DE PANICO']})]
my_dictSNP=readSNPfile('dataSNP.txt') 
writeSNPfile(my_dictSNP,'dataSNPfinal.txt')
my_dictSNP_copy=readSNPfile('dataSNPfinal.txt') 
for (codeRS, solution) in tests:
    my_dictSNP_nuevo=readSNPfile('dataSNP.txt') 
    print('  codigoRS: {}\t resultado: {}'.format(codeRS, solution))
    resultado=my_dictSNP_copy.get(codeRS, {})
    try:     
        assert equalitySNP(resultado, solution) and equalitySNPs(my_dictSNP_nuevo, my_dictSNP_copy, [])
    except:
        print('    RESULTADO INCORRECTO codigoRS: {} resultado: {}'.format(codeRS, resultado))
print('FIN de las pruebas (writeSNPfile())')
### END TESTS
```