# Programación I. Segunda práctica obligatoria
En el **Nim**, los jugadores colocan una cantidad arbitraria de fichas sobre la superficie del juego, separadas en grupos. Tanto la cantidad de grupos como la cantidad de fichas en cada grupo son también arbitrarias. En cada turno, el jugador correspondiente quita cualquier cantidad de fichas de un solo grupo (entre uno y la cantidad de fichas del grupo), y los jugadores se alternan retirando fichas. La partida finaliza cuando un jugador quita la última ficha de la superficie del juego, declarándose entonces ganador. El propósito de la práctica es codificar algunas funciones del juego. Para ello, asumimos que la cantidad de grupos es diez y la cantidad máxima de fichas en cada grupo es 20.

### Ejercicio 1
Diseña una función llamada `initGame()` que prepare una partida, insertando los valores introducidos por el teclado en una lista de enteros de tamaño adecuado. La función debe comprobar que los datos insertados en la lista son válidos.

<!-- grade_id: initGame_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import math
CANTIDAD_MONTONES = 10
CANTIDAD_PIEZAS = 20
def getPieces():
    dato = 0
    while (dato<=0 or dato>=CANTIDAD_PIEZAS+1):
        try:
            dato = int(input(" Introduce una cantidad valida de piezas (0<=cantidad<=20): "))
        except ValueError:
            print(" Debes introducir un valor de tipo entero")
            dato = 0
    return dato
def initGame():
    listaJuego=[]
    for posicion in range(0, CANTIDAD_MONTONES, 1):
        # print("Cantidad de piezas del grupo "+str(posicion), end = " ")
        listaJuego.append(getPieces())
    return listaJuego
### END SOLUTION
```

<!-- grade_id: initGame_test, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
import unittest
from unittest.mock import patch
class InitTestClass(unittest.TestCase):
    test_init = (['-1','0','23','9','1','0','18','21','0','7','13','9','11','-3','4','2','20'],[9,1,18,7,13,9,11,4,2,20])
    @patch('builtins.input', side_effect = test_init[0])
    def test_initGame(self, input):
        self.assertEqual(initGame(), InitTestClass.test_init[1]) 
print('INICIO de las pruebas (initGame())')
a = InitTestClass()
print("Entrada: "+str(InitTestClass.test_init[0]))
print("Respuesta: "+str(InitTestClass.test_init[1]))
try:
    a.test_initGame()
except AssertionError:
    print("  Prueba incorrecta")
print('FIN de las pruebas (initGame())')
### END TESTS
```

Output:
```
INICIO de las pruebas (initGame())
Entrada: ['-1', '0', '23', '9', '1', '0', '18', '21', '0', '7', '13', '9', '11', '-3', '4', '2', '20']
Respuesta: [9, 1, 18, 7, 13, 9, 11, 4, 2, 20]
FIN de las pruebas (initGame())
```

### Ejercicio 2
Diseña una función llamada `deletePiecesPos()` que, dadas la lista `lst` de una partida, una posición `p` y una cantidad de fichas `c`, elimine la cantidad de fichas `c` de la posición `p` de la lista `lst`. La función debe comprobar que tanto la posición `p` como la cantidad de fichas `c` son válidas.

<!-- grade_id: deletePiecesPos_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def deletePiecesPos(lst,p,c):
    if (p<0 or p>len(lst)-1) or (c<=0 or c>=lst[p]+1):
        raise ValueError("Parametros incorrectos")
    lst[p]=lst[p]-c
### END SOLUTION
```

<!-- grade_id: deletePiecesPos_test, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
print('INICIO de las pruebas (deletePiecesPos())')
tests = [([20,12,3,4,8,0,11,5,1,13],0,5,[15,12,3,4,8,0,11,5,1,13]),
         ([20,12,3,4,8,0,11,5,1,3],2,1,[20,12,2,4,8,0,11,5,1,3]),
         ([20,12,2,4,8,0,11,5,1,3],2,2,[20,12,0,4,8,0,11,5,1,3]),
         ([20,12,2,4,8,0,11,5,1,3],9,3,[20,12,2,4,8,0,11,5,1,0])]
for (lst, p, c, resultado) in tests:
    print('   lista: {} posicion: {} cantidad: {} resultado: {}'.format(lst, p, c, resultado))
    deletePiecesPos(lst, p, c)
    try:
        assert resultado==lst
    except:
        print('      Resultado incorrecto: '+str(lst))
print('FIN de las pruebas (deletePiecesPos())')
### END TESTS
```

Output:
```
INICIO de las pruebas (deletePiecesPos())
   lista: [20, 12, 3, 4, 8, 0, 11, 5, 1, 13] posicion: 0 cantidad: 5 resultado: [15, 12, 3, 4, 8, 0, 11, 5, 1, 13]
   lista: [20, 12, 3, 4, 8, 0, 11, 5, 1, 3] posicion: 2 cantidad: 1 resultado: [20, 12, 2, 4, 8, 0, 11, 5, 1, 3]
   lista: [20, 12, 2, 4, 8, 0, 11, 5, 1, 3] posicion: 2 cantidad: 2 resultado: [20, 12, 0, 4, 8, 0, 11, 5, 1, 3]
   lista: [20, 12, 2, 4, 8, 0, 11, 5, 1, 3] posicion: 9 cantidad: 3 resultado: [20, 12, 2, 4, 8, 0, 11, 5, 1, 0]
FIN de las pruebas (deletePiecesPos())
```

### Ejercicio 3
Diseña una función llamada `endGame()` que decida si una partida dada ha terminado, es decir, si su lista `lst` asociada contiene únicamente valores nulos.

<!-- grade_id: endGame_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def endGame(lst):
    posicion = 0
    while (posicion<len(lst)) and (lst[posicion]==0):
        posicion = posicion+1
    return posicion==len(lst)
### END SOLUTION
```

<!-- grade_id: endGame_test, points: 1, grade: true, solution: false -->
```python
print('INICIO de las pruebas (endGame())')
tests = [([0,0,0,0,0,0,0,0,0,0],True),
         ([0,0,0,0,5,0,0,0,0,0],False),
         ([0,0,0,0,0,0,0,0,0,17],False),
         ([0,0,19,13,0,0,0,11,0,0],False)]
for (lst, resultado) in tests:
    print('   lista: {} resultado: {}'.format(lst, resultado))
    result = endGame(lst)
    try:
        assert resultado==result
    except:
        print('      Resultado incorrecto: '+str(result))
print('FIN de las pruebas (endGame())')
```

Output:
```
INICIO de las pruebas (endGame())
   lista: [0, 0, 0, 0, 0, 0, 0, 0, 0, 0] resultado: True
   lista: [0, 0, 0, 0, 5, 0, 0, 0, 0, 0] resultado: False
   lista: [0, 0, 0, 0, 0, 0, 0, 0, 0, 17] resultado: False
   lista: [0, 0, 19, 13, 0, 0, 0, 11, 0, 0] resultado: False
FIN de las pruebas (endGame())
```

### Ejercicio 4a
Diseña una función llamada isDenierForm() que determine si todos los grupos de una partida dada tienen una única ficha excepto uno, que tendrá varias piezas. En caso afirmativo, la función devolverá el grupo que tiene tal cantidad de fichas y en caso contrario, la función devolverá −1.

<!-- grade_id: isDenierForm_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def isDenierForm(lst):
    contadorNoUnos = 0
    posicion = 0
    parar = False
    while (posicion<len(lst)) and (not parar) and (contadorNoUnos<=1):
        if (lst[posicion]==0):
            parar = True
        elif (lst[posicion]!=1):
            contadorNoUnos = contadorNoUnos+1
            posicionNoUno = posicion
        posicion = posicion+1
    resultado = -1
    if ((contadorNoUnos==1) and (lst[posicionNoUno]>1)):
        resultado = posicionNoUno
    return resultado
### END SOLUTION
```

<!-- grade_id: isDenierForm_test, points: 2, grade: true, solution: false -->
```python
print('INICIO de las pruebas (isDenierForm())')
tests = [([1,1,1,12,1,1,1,1,1,1],3),
         ([1,1,1,0,1,1,1,1,1,1],-1),
         ([12,1,1,1,1,1,1,1,1,1],0),
         ([1,1,1,12,1,1,1,1,1,11],-1),
         ([1,1,1,13,19,1,1,1,1,1],-1),
         ([1,1,1,1,1,1,1,1,1,12],9),
         ([1,1,0,12,1,1,1,1,1,11],-1),
         ([1,1,1,1,1,1,1,1,1,1],-1)]
for (lst, resultado) in tests:
    print('   lista: {} resultado: {}'.format(lst, resultado))
    result = isDenierForm(lst)
    try:
        assert resultado==result
    except:
        print('      Resultado incorrecto: '+str(result))
print('FIN de las pruebas (isDenierForm())')
```

Output:
```
INICIO de las pruebas (isDenierForm())
   lista: [1, 1, 1, 12, 1, 1, 1, 1, 1, 1] resultado: 3
   lista: [1, 1, 1, 0, 1, 1, 1, 1, 1, 1] resultado: -1
   lista: [12, 1, 1, 1, 1, 1, 1, 1, 1, 1] resultado: 0
   lista: [1, 1, 1, 12, 1, 1, 1, 1, 1, 11] resultado: -1
   lista: [1, 1, 1, 13, 19, 1, 1, 1, 1, 1] resultado: -1
   lista: [1, 1, 1, 1, 1, 1, 1, 1, 1, 12] resultado: 9
   lista: [1, 1, 0, 12, 1, 1, 1, 1, 1, 11] resultado: -1
   lista: [1, 1, 1, 1, 1, 1, 1, 1, 1, 1] resultado: -1
FIN de las pruebas (isDenierForm())
```

### Ejercicio 4b
Diseña una función llamada `maxPiecesGroup()` que determine el grupo de una partida dada que tiene la mayor cantidad de fichas. Si la partida tiene varios grupos con tal cantidad, entonces la función devolverá el primero de ellos. 
Diseña una función llamada `saturateGame()` que, dados la lista `lst` de una partida y un entero `e`, sature la lista con el entero, es decir, asigne a cada posición `p` de la lista `lst` el valor `max(e,lst[p])`.

<!-- grade_id: maxPiecesGroup_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def maxPiecesGroup(lst):
    maximo = -math.inf
    for posicion in range(0,len(lst),1):
        if (lst[posicion]==CANTIDAD_PIEZAS):
            posicionMaximo = posicion
            break
        if (maximo<lst[posicion]):
            maximo = lst[posicion]
            posicionMaximo = posicion
    return posicionMaximo
def saturateGame(lst,e):
    for posicion in range(0,len(lst),1):
        lst[posicion] = max(lst[posicion],e)
### END SOLUTION
```

<!-- grade_id: maxPiecesGroup_test, points: 2, grade: true, solution: false -->
```python
print('INICIO de las pruebas (maxPiecesGroup())')
tests = [([20,12,3,4,8,0,11,5,1,13],0),
         ([11,12,3,4,8,0,11,5,1,20],9),
         ([1,1,1,12,1,1,1,12,1,1],3),
         ([1,7,1,12,9,1,5,7,1,4],3)]
for (lst, resultado) in tests:
    print('   lista: {} resultado: {}'.format(lst, resultado))
    result = maxPiecesGroup(lst)
    try:
        assert resultado==result
    except:
        print('      Resultado incorrecto: '+str(result))
print('FIN de las pruebas (maxPiecesGroup())')
print('INICIO de las pruebas (saturateGame())')
tests = [([20,12,3,4,8,2,11,5,1,13],1,[20,12,3,4,8,2,11,5,1,13]),
         ([11,12,3,4,8,0,11,5,1,20],13,[13,13,13,13,13,13,13,13,13,20]),
         ([1,1,1,12,1,1,1,12,1,1],3,[3,3,3,12,3,3,3,12,3,3]),
         ([1,7,1,12,9,1,5,7,1,4],3,[3,7,3,12,9,3,5,7,3,4])]
for (lst, e, resultado) in tests:
    print('   lista: {} e: {} resultado: {}'.format(lst, e, resultado))
    saturateGame(lst, e)
    try:
        assert resultado==lst
    except:
        print('      Resultado incorrecto: '+str(lst))
print('FIN de las pruebas (saturateGame())')
```

Output:
```
INICIO de las pruebas (maxPiecesGroup())
   lista: [20, 12, 3, 4, 8, 0, 11, 5, 1, 13] resultado: 0
   lista: [11, 12, 3, 4, 8, 0, 11, 5, 1, 20] resultado: 9
   lista: [1, 1, 1, 12, 1, 1, 1, 12, 1, 1] resultado: 3
   lista: [1, 7, 1, 12, 9, 1, 5, 7, 1, 4] resultado: 3
FIN de las pruebas (maxPiecesGroup())
INICIO de las pruebas (saturateGame())
   lista: [20, 12, 3, 4, 8, 2, 11, 5, 1, 13] e: 1 resultado: [20, 12, 3, 4, 8, 2, 11, 5, 1, 13]
   lista: [11, 12, 3, 4, 8, 0, 11, 5, 1, 20] e: 13 resultado: [13, 13, 13, 13, 13, 13, 13, 13, 13, 20]
   lista: [1, 1, 1, 12, 1, 1, 1, 12, 1, 1] e: 3 resultado: [3, 3, 3, 12, 3, 3, 3, 12, 3, 3]
   lista: [1, 7, 1, 12, 9, 1, 5, 7, 1, 4] e: 3 resultado: [3, 7, 3, 12, 9, 3, 5, 7, 3, 4]
FIN de las pruebas (saturateGame())
```

### Ejercicio 4c
Diseña una función llamada `deletePiecesDenier()` que aplique la estrategia siguiente: 
* Si la partida es Denier y la cantidad de fichas del grupo con varias piezas es impar, entonces la función eliminará la mitad más uno de la fichas de ese grupo. Si la partida es Denier y la cantidad de fichas del montón con varias piezas es par, entonces la función eliminará la mitad de la fichas de ese montón. 
* Si la partida no es Denier, entonces la función eliminará una ficha del grupo con la mayor cantidad de fichas.

<!-- grade_id: deletePiecesDenier_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def deletePiecesDenier(lst):
    grupoNoUnos = isDenierForm(lst)
    if (grupoNoUnos!=-1):
        if (lst[grupoNoUnos]%2==0):
            deletePiecesPos(lst, grupoNoUnos, lst[grupoNoUnos]//2)
        else:    
            deletePiecesPos(lst, grupoNoUnos, lst[grupoNoUnos]//2+1)
    else:
        grupoMaximo = maxPiecesGroup(lst)
        deletePiecesPos(lst, grupoMaximo, 1)
### END SOLUTION
```

<!-- grade_id: deletePiecesDenier_test, points: 2, grade: true, solution: false -->
```python
print('INICIO de las pruebas (deletePiecesDenier())')
tests = [([1,1,1,12,1,1,1,1,1,1],[1,1,1,6,1,1,1,1,1,1]),
         ([20,12,3,4,8,0,11,5,1,13],[19,12,3,4,8,0,11,5,1,13]), 
         ([1,1,1,1,1,1,1,1,1,13],[1,1,1,1,1,1,1,1,1,6]),
         ([1,1,1,12,1,1,1,12,1,1],[1,1,1,11,1,1,1,12,1,1])]
for (lst, resultado) in tests:
    print('   lista: {} resultado: {}'.format(lst, resultado))
    deletePiecesDenier(lst)
    try:
        assert resultado==lst
    except:
        print('      Resultado incorrecto: '+str(lst))
print('FIN de las pruebas (deletePiecesDenier())')
```

Output:
```
INICIO de las pruebas (deletePiecesDenier())
   lista: [1, 1, 1, 12, 1, 1, 1, 1, 1, 1] resultado: [1, 1, 1, 6, 1, 1, 1, 1, 1, 1]
   lista: [20, 12, 3, 4, 8, 0, 11, 5, 1, 13] resultado: [19, 12, 3, 4, 8, 0, 11, 5, 1, 13]
   lista: [1, 1, 1, 1, 1, 1, 1, 1, 1, 13] resultado: [1, 1, 1, 1, 1, 1, 1, 1, 1, 6]
   lista: [1, 1, 1, 12, 1, 1, 1, 12, 1, 1] resultado: [1, 1, 1, 11, 1, 1, 1, 12, 1, 1]
FIN de las pruebas (deletePiecesDenier())
```

### Ejercicio 5
Diseña una función llamada `playGame()` que prepare una partida y efectúe una cantidad de jugadas dada empleando la estrategia Denier. La partida debe finalizar tan pronto como se hayan eliminado todas las fichas de todos los grupos o se haya alcanzado la cantidad de jugadas dada.

<!-- grade_id: playGame_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def playGame(contador):
    lst = initGame() 
    cantidadTiradas = 0       
    # print("Reparto inicial de las fichas: "+str(lst))
    while ((not endGame(lst)) and (cantidadTiradas<abs(contador))):
        deletePiecesDenier(lst)
        # print("Estado del juego (jugada " + str(cantidadTiradas+1) + "): " +str(lst))
        cantidadTiradas = cantidadTiradas+1
    return lst
### END SOLUTION
```

<!-- grade_id: playGame_test, points: 1, grade: true, solution: false -->
```python
import unittest
from unittest.mock import patch
class PlayTestClass(unittest.TestCase):
    test_init = (['1','1','1','12','1','1','1','1','1','1'],10,[0,0,0,0,0,0,0,1,1,1])
    @patch('builtins.input', side_effect = test_init[0])
    def test_playGame(self, input):
        self.assertEqual(playGame(PlayTestClass.test_init[1]), PlayTestClass.test_init[2]) 
print('INICIO de las pruebas (playGame())')
a = PlayTestClass()
print("Entrada: "+str(PlayTestClass.test_init[0]))
print("Cantidad de jugadas: "+str(PlayTestClass.test_init[1]))
print("Respuesta: "+str(PlayTestClass.test_init[2]))
try:
    a.test_playGame()
except AssertionError:
    print("  Prueba incorrecta")
print('FIN de las pruebas (playGame())')
```

Output:
```
INICIO de las pruebas (playGame())
Entrada: ['1', '1', '1', '12', '1', '1', '1', '1', '1', '1']
Cantidad de jugadas: 10
Respuesta: [0, 0, 0, 0, 0, 0, 0, 1, 1, 1]
FIN de las pruebas (playGame())
```