# Programación I. Primera práctica obligatoria

Ejecuta la celda siguiente para realizar el testeo de tus funciones.

```python
def test_movement(piece,test_list):
    for (f_ori,c_ori,f_dest,c_dest,res) in tests:
        print('origen: ({},{}) destino: ({},{}) resultado:{}'.format(f_ori, c_ori, f_dest, c_dest, res))
        result=piece(f_ori, c_ori, f_dest, c_dest)
        try:
            assert res==result
        except AssertionError:
            print('  Resultado incorrecto: '+str(result))
```

### Ejercicio 1 **[Jaque mate]**
Diseña y escribe una función que, dadas dos posiciones del tablero de ajedrez, decida si la figura del Rey puede moverse desde la primera posición hasta la segunda. Los datos de entrada de la función son la posición inicial de la figura `(r_ori,c_ori)` y su posición final `(r_dest,c_dest)`, donde las filas son enteros y las columnas son letras tales que `1≤r_ori,r_dest≤8` y `'a'≤c_ori,c_dest≤'h'`. El dato de salida será un valor booleano que indica si el movimiento es o no posible. El Rey puede moverse en cualquier dirección (vertical, horizontal y diagonales) avanzando o retrocediendo una sola casilla.  
El nombre de la función y el orden de los parámetros de entrada son los siguientes: `king_movement(r_ori, c_ori, r_dest, c_dest)`  
**Observación:** No es necesario comprobar la corrección de los datos de entrada.

<!-- grade_id: function_king_movement, points: 0, grade: false, solution: true -->
```python
def king_movement(r_ori,c_ori,r_dest,c_dest):
    ### BEGIN SOLUTION
    filaOrigen,columnaOrigen,filaDestino,columnaDestino=r_ori,c_ori,r_dest,c_dest
    columnaOrigen_int=ord(columnaOrigen)-ord('a')+1
    columnaDestino_int=ord(columnaDestino)-ord('a')+1
    if (filaOrigen==filaDestino) and (abs(columnaOrigen_int-columnaDestino_int)==1):
        return True
    elif (abs(filaOrigen-filaDestino)==1) and (columnaOrigen_int==columnaDestino_int):
        return True
    elif (abs(filaOrigen-filaDestino)==1) and (abs(columnaOrigen_int-columnaDestino_int)==1):
        return True
    else:
        return False
    ### END SOLUTION
```

<!-- grade_id: test_king_movement, points: 1, grade: true, solution: false -->
```python
tests = [(3,'f',4,'e',True),(3,'f',3,'e',True),(3,'f',2,'e',True),
         (3,'f',4,'f',True),(3,'f',2,'f',True),(3,'f',4,'g',True),
         (3,'f',3,'g',True),(3,'f',2,'g',True),(8,'a',8,'c',False),
         (8,'a',6,'d',False),(8,'a',4,'a',False),(3,'f',3,'f',False)]
test_movement(king_movement, tests)
```

Output:
```
origen: (3,f) destino: (4,e) resultado:True
origen: (3,f) destino: (3,e) resultado:True
origen: (3,f) destino: (2,e) resultado:True
origen: (3,f) destino: (4,f) resultado:True
origen: (3,f) destino: (2,f) resultado:True
origen: (3,f) destino: (4,g) resultado:True
origen: (3,f) destino: (3,g) resultado:True
origen: (3,f) destino: (2,g) resultado:True
origen: (8,a) destino: (8,c) resultado:False
origen: (8,a) destino: (6,d) resultado:False
origen: (8,a) destino: (4,a) resultado:False
origen: (3,f) destino: (3,f) resultado:False
```

### Ejercicio 2 **[Otros trebejos]** 
Diseña y escribe funciones similares a `king_movement()` para estudiar los movimientos de la Reina (`queen_movement()`), la Torre (`rook_movement()`), el Alfil (`bishop_movement()`) y el Caballo (`knight_movement()`). Dadas dos posiciones del tablero de ajedrez, cada función debe decidir si la figura asociada puede moverse desde la primera posición hasta la segunda. Los datos de entrada de cada función son la posición inicial de la figura  `(r_ori,c_ori)` y su posición final `(r_dest,c_dest)`, donde las filas son enteros y las columnas son letras tales que `1≤r_ori,r_dest≤8` y `'a'≤c_ori,c_dest≤'h'`. El dato de salida será un valor booleano que indica si el movimiento es o no posible.  
**Observación:** No es necesario comprobar la corrección de los datos de entrada.

<!-- grade_id: function_queen_movement, points: 0, grade: false, solution: true -->
```python
def queen_movement(r_ori,c_ori,r_dest,c_dest):
    ### BEGIN SOLUTION
    filaOrigen,columnaOrigen,filaDestino,columnaDestino=r_ori,c_ori,r_dest,c_dest
    columnaOrigen_int=ord(columnaOrigen)-ord('a')+1
    columnaDestino_int=ord(columnaDestino)-ord('a')+1
    if (filaOrigen==filaDestino) and (columnaOrigen_int==columnaDestino_int):
        return False
    elif (filaOrigen==filaDestino):
        return True
    elif (columnaOrigen_int==columnaDestino_int):
        return True
    elif (abs(filaOrigen-filaDestino)==abs(columnaOrigen_int-columnaDestino_int)):
        return True
    else:
        return False
    ### END SOLUTION
```

<!-- grade_id: test_queen_movement, points: 1, grade: true, solution: false -->
```python
tests = [(1,'a',7,'a',True),(1,'a',8,'h',True),(1,'a',1,'b',True),(1,'a',2,'c',False),
         (4,'h',1,'e',True),(4,'h',4,'d',True),(4,'h',8,'d',True),(4,'h',8,'e',False),
         (8,'d',5,'a',True),(8,'d',5,'g',True),(8,'d',7,'d',True),(8,'d',6,'c',False),
         (3,'c',5,'a',True),(3,'c',3,'d',True),(3,'c',5,'e',True),(3,'c',3,'c',False)]
test_movement(queen_movement, tests)
```

Output:
```
origen: (1,a) destino: (7,a) resultado:True
origen: (1,a) destino: (8,h) resultado:True
origen: (1,a) destino: (1,b) resultado:True
origen: (1,a) destino: (2,c) resultado:False
origen: (4,h) destino: (1,e) resultado:True
origen: (4,h) destino: (4,d) resultado:True
origen: (4,h) destino: (8,d) resultado:True
origen: (4,h) destino: (8,e) resultado:False
origen: (8,d) destino: (5,a) resultado:True
origen: (8,d) destino: (5,g) resultado:True
origen: (8,d) destino: (7,d) resultado:True
origen: (8,d) destino: (6,c) resultado:False
origen: (3,c) destino: (5,a) resultado:True
origen: (3,c) destino: (3,d) resultado:True
origen: (3,c) destino: (5,e) resultado:True
origen: (3,c) destino: (3,c) resultado:False
```

<!-- grade_id: function_rook_movement, points: 0, grade: false, solution: true -->
```python
def rook_movement(r_ori,c_ori,r_dest,c_dest):
    ### BEGIN SOLUTION
    filaOrigen,columnaOrigen,filaDestino,columnaDestino=r_ori,c_ori,r_dest,c_dest
    columnaOrigen_int=ord(columnaOrigen)-ord('a')+1
    columnaDestino_int=ord(columnaDestino)-ord('a')+1
    if (filaOrigen==filaDestino) and (columnaOrigen_int==columnaDestino_int):
        return False
    elif (filaOrigen==filaDestino):
        return True
    elif (columnaOrigen_int==columnaDestino_int):
        return True
    else:
        return False 
    ### END SOLUTION
```

<!-- grade_id: test_rook_movement, points: 1, grade: true, solution: false -->
```python
tests = [(5,'d',2,'d',True),(5,'d',7,'d',True),(5,'d',5,'e',True),(5,'d',5,'a',True),
         (7,'g',8,'f',False),(7,'g',6,'h',False),(7,'g',5,'e',False),(7,'g',7,'g',False)]
test_movement(rook_movement, tests)
```

Output:
```
origen: (5,d) destino: (2,d) resultado:True
origen: (5,d) destino: (7,d) resultado:True
origen: (5,d) destino: (5,e) resultado:True
origen: (5,d) destino: (5,a) resultado:True
origen: (7,g) destino: (8,f) resultado:False
origen: (7,g) destino: (6,h) resultado:False
origen: (7,g) destino: (5,e) resultado:False
origen: (7,g) destino: (7,g) resultado:False
```

<!-- grade_id: function_bishop_movement, points: 0, grade: false, solution: true -->
```python
def bishop_movement(r_ori,c_ori,r_dest,c_dest):
    ### BEGIN SOLUTION
    filaOrigen,columnaOrigen,filaDestino,columnaDestino=r_ori,c_ori,r_dest,c_dest
    columnaOrigen_int=ord(columnaOrigen)-ord('a')+1
    columnaDestino_int=ord(columnaDestino)-ord('a')+1
    if (filaOrigen==filaDestino) and (columnaOrigen_int==columnaDestino_int):
        return False
    elif (abs(filaOrigen-filaDestino)==abs(columnaOrigen_int-columnaDestino_int)):
        return True
    else:
        return False
    ### END SOLUTION
```

<!-- grade_id: test_bishop_movement, points: 1, grade: true, solution: false -->
```python
tests = [(1,'f',3,'d',True),(1,'f',2,'g',True),(2,'g',3,'h',True),(2,'g',8,'a',True),
         (8,'f',6,'h',True),(8,'f',3,'a',True),(6,'h',8,'g',False),(6,'h',7,'h',False),
         (4,'d',6,'e',False),(4,'d',1,'d',False),(2,'b',2,'h',False),(2,'b',2,'b',False)]
test_movement(bishop_movement, tests)
```

Output:
```
origen: (1,f) destino: (3,d) resultado:True
origen: (1,f) destino: (2,g) resultado:True
origen: (2,g) destino: (3,h) resultado:True
origen: (2,g) destino: (8,a) resultado:True
origen: (8,f) destino: (6,h) resultado:True
origen: (8,f) destino: (3,a) resultado:True
origen: (6,h) destino: (8,g) resultado:False
origen: (6,h) destino: (7,h) resultado:False
origen: (4,d) destino: (6,e) resultado:False
origen: (4,d) destino: (1,d) resultado:False
origen: (2,b) destino: (2,h) resultado:False
origen: (2,b) destino: (2,b) resultado:False
```

<!-- grade_id: function_knight_movement, points: 0, grade: false, solution: true -->
```python
def knight_movement(r_ori,c_ori,r_dest,c_dest):
    ### BEGIN SOLUTION
    filaOrigen,columnaOrigen,filaDestino,columnaDestino=r_ori,c_ori,r_dest,c_dest
    columnaOrigen_int=ord(columnaOrigen)-ord('a')+1
    columnaDestino_int=ord(columnaDestino)-ord('a')+1
    if (columnaDestino_int-columnaOrigen_int==2) and (abs(filaOrigen-filaDestino)==1):
        return True
    elif (columnaDestino_int-columnaOrigen_int==-2) and (abs(filaOrigen-filaDestino)==1):
        return True
    elif (filaDestino-filaOrigen==2) and (abs(columnaOrigen_int-columnaDestino_int)==1):
        return True
    elif (filaDestino-filaOrigen==-2) and (abs(columnaOrigen_int-columnaDestino_int)==1):
        return True
    else:
        return False
    ### END SOLUTION
```

<!-- grade_id: test_knight_movement, points: 1, grade: true, solution: false -->
```python
tests = [(4,'e',6,'d',True),(4,'e',6,'f',True),(4,'e',5,'c',True),(4,'e',5,'g',True),
         (4,'e',3,'c',True),(4,'e',3,'g',True),(4,'e',2,'d',True),(4,'e',2,'f',True),
         (8,'b',8,'c',False),(8,'b',5,'e',False),(8,'b',6,'b',False),(8,'b',8,'b',False)]
test_movement(knight_movement, tests)
```

Output:
```
origen: (4,e) destino: (6,d) resultado:True
origen: (4,e) destino: (6,f) resultado:True
origen: (4,e) destino: (5,c) resultado:True
origen: (4,e) destino: (5,g) resultado:True
origen: (4,e) destino: (3,c) resultado:True
origen: (4,e) destino: (3,g) resultado:True
origen: (4,e) destino: (2,d) resultado:True
origen: (4,e) destino: (2,f) resultado:True
origen: (8,b) destino: (8,c) resultado:False
origen: (8,b) destino: (5,e) resultado:False
origen: (8,b) destino: (6,b) resultado:False
origen: (8,b) destino: (8,b) resultado:False
```

### Ejercicio 3 **[Ajedrez]**
Diseña y escribe una función que, dadas una pieza de ajedrez y dos posiciones del tablero, decida si dicha pieza puede moverse desde la primera posición hasta la segunda. Los datos de entrada de la función son la pieza `p`, su posición inicial `(r_ori,c_ori)` y la posición final `(r_dest,c_dest)`. La pieza `p` es un carácter representando una figura ('K': Rey, 'Q': Reina, 'R': Torre, 'B': Alfil, 'N': Caballo), las filas son enteros y las columnas son letras tales que `1≤r_ori,r_dest≤8` y `'a'≤c_ori,c_dest≤'h'`. El dato de salida será un valor booleano que indica si el movimiento es o no posible.  
El nombre de la función y el orden de los parámetros de entrada son los siguientes: `chess_movement(p, r_ori, c_ori, r_dest, c_dest)`  
**Observación:** No es necesario comprobar la corrección de los datos de entrada.

<!-- grade_id: function_chess_movement, points: 0, grade: false, solution: true -->
```python
def chess_movement(p,r_ori,c_ori,r_dest,c_dest):
    ### BEGIN SOLUTION
    pieza,filaOrigen,columnaOrigen,filaDestino,columnaDestino=p,r_ori,c_ori,r_dest,c_dest
    if (pieza=='K'):
        return king_movement(filaOrigen,columnaOrigen,filaDestino,columnaDestino)
    elif (pieza=='Q'):
        return queen_movement(filaOrigen,columnaOrigen,filaDestino,columnaDestino)
    elif (pieza=='N'):
        return knight_movement(filaOrigen,columnaOrigen,filaDestino,columnaDestino)
    elif (pieza=='R'):
        return rook_movement(filaOrigen,columnaOrigen,filaDestino,columnaDestino)
    elif (pieza=='B'):
        return bishop_movement(filaOrigen,columnaOrigen,filaDestino,columnaDestino)
    else:
        return False
    ### END SOLUTION
```

<!-- grade_id: test_chess_movement, points: 1, grade: true, solution: false -->
```python
tests = [('K',8,'b',8,'c',True),('B',1,'f',3,'d',True),('R',7,'g',8,'f',False),('K',3,'f',2,'f',True),('N',8,'b',8,'c',False),
         ('N',4,'e',6,'d',True),('B',6,'h',7,'h',False),('Q',1,'a',7,'a',True),('Q',1,'a',2,'c',False),('R',5,'d',2,'d',True)]
for (p,f_ori,c_ori,f_dest,c_dest,res) in tests:
    print('origen: ({},{}) destino: ({},{}) resultado:{}'.format(f_ori, c_ori, f_dest, c_dest, res))
    result=chess_movement(p, f_ori, c_ori, f_dest, c_dest)
    try:
        assert res==result
    except AssertionError:
        print('  Resultado incorrecto: '+str(result))
```

Output:
```
origen: (8,b) destino: (8,c) resultado:True
origen: (1,f) destino: (3,d) resultado:True
origen: (7,g) destino: (8,f) resultado:False
origen: (3,f) destino: (2,f) resultado:True
origen: (8,b) destino: (8,c) resultado:False
origen: (4,e) destino: (6,d) resultado:True
origen: (6,h) destino: (7,h) resultado:False
origen: (1,a) destino: (7,a) resultado:True
origen: (1,a) destino: (2,c) resultado:False
origen: (5,d) destino: (2,d) resultado:True
```