# Programación II. Primera práctica obligatoria
Podemos representar un polinomio de grado $n$, $p(x)\!=\!a_nx^n\!+\!a_{n-1}x^{n-1}\!+\!...\!+\!a_1x\!+\!a_0$, mediante la lista de sus coeficientes $[a_n,a_{n-1},...,a_1,a_0]$. Por ejemplo, $p_1(x)\!=\!x^3\!+\!2x^2\!+\!x\!+\!1$ puede representarse mediante la lista $[1,2,1,1]$, y podemos representar el polinomio $p_2(x)\!=\!x^3\!+\!2x^2\!+\!x$ mediante la lista $[1,2,1,0]$. Observa que la representación de $p_3(x)\!=\!3x^2\!+\!2x\!+\!1$ es la lista $[3,2,1]$, y que la lista correspondiente al polinomio $p_4(x)\!=\!-x^3$ es $[-1,0,0,0]$.

### Ejercicio 1
Diseña una función llamada `polynomial_isSolution(p,n)` que, dados un polinomio `p` y un número entero `n`, decida si `n` es o no un cero de `p` (los ceros de un polinomio $p(x)$ son los valores de $x$ para los que $p(x)\!=\!0$). Por ejemplo, el resultado de la llamada `polynomial_isSolution([1,0,-1],1)` debe ser `True`, ya que el polinomio $p(x)\!=\!x^{2}\!-\!1$ se anula en $x\!=\!1$ (y en $x\!=\!-1$).

<!-- grade_id: isSolution_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def polynomial_isSolution(p,n):
    valor=0
    for i in range(-1,-len(p)-1,-1):
        valor=valor+p[i]*(n**(-i-1))
    return valor==0
### END SOLUTION
```

<!-- grade_id: isSolution_test, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
print('INICIO de las pruebas (polynomial_isSolution())')
tests = [([],7,True),([-3],0,False),([-3],3,False),
         ([1,0,-1],-1,True),([1,0,-1],0,False),([1,0,-1],1,True),
         ([1,2,-8,-18,-9],-3,True),([1,2,-8,-18,-9],-1,True),
         ([1,2,-8,-18,-9],0,False),([1,2,-8,-18,-9],3,True)]
for (p, n, resultado) in tests:
    print('   polinomio: {} valor: {} resultado: {}'.format(p, n, resultado))
    p_copia=p.copy()
    result=polynomial_isSolution(p_copia, n)
    try:
        assert resultado==result and p==p_copia
    except:
        print('      RESULTADO INCORRECTO / MODIFICACIÓN DE PARÁMETROS - polinomio: {} resultado: {}'.format(p_copia, result))
print('FIN de las pruebas (polynomial_isSolution())')
### END TESTS
```

Output:
```
INICIO de las pruebas (polynomial_isSolution())
   polinomio: [] valor: 7 resultado: True
   polinomio: [-3] valor: 0 resultado: False
   polinomio: [-3] valor: 3 resultado: False
   polinomio: [1, 0, -1] valor: -1 resultado: True
   polinomio: [1, 0, -1] valor: 0 resultado: False
   polinomio: [1, 0, -1] valor: 1 resultado: True
   polinomio: [1, 2, -8, -18, -9] valor: -3 resultado: True
   polinomio: [1, 2, -8, -18, -9] valor: -1 resultado: True
   polinomio: [1, 2, -8, -18, -9] valor: 0 resultado: False
   polinomio: [1, 2, -8, -18, -9] valor: 3 resultado: True
FIN de las pruebas (polynomial_isSolution())
```

### Ejercicio 2
Diseña una función llamada `polynomial_sum(p1,p2)` que, dados dos polinomios `p1` y `p2`, calcule la suma `p1+p2`. Por ejemplo, el resultado de la llamada `polynomial_sum([1,2,1,1],[3,2,1])` debe ser la lista de coeficientes correspondiente al polinomio $q(x)\!=\!x^3\!+\!5x^2\!+\!3x\!+\!2$.

<!-- grade_id: sum_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def polynomial_normalization(polinomio):
    posicion=0
    while (posicion<len(polinomio) and polinomio[posicion]==0):
        polinomio.pop(0)
def coeficiente(polinomio,pos): # pos<0
    if (len(polinomio)+pos>=0):
        return polinomio[pos]
    else:
        return 0
def polynomial_sum(p1,p2):
    suma=[0]*max(len(p1),len(p2))
    for i in range(-1,-len(suma)-1,-1):
        suma[i]=coeficiente(p1,i)+coeficiente(p2,i)
    polynomial_normalization(suma)
    return suma
### END SOLUTION
```

<!-- grade_id: sum_test, points: 3, grade: true, solution: false -->
```python
### BEGIN TESTS
print('INICIO de las pruebas (polynomial_sum())')
tests = [([],[3],[3]),([-8],[-5,6,-2,3,8],[-5,6,-2,3,0]),
         ([1,1],[],[1,1]),([1,0,-1],[-1,0,1],[]),
         ([1,5,0,2],[-3,0],[1,5,-3,2]),
         ([1,2,1,1],[3,2,1],[1,5,3,2]),
         ([-5,6,-2,3,8],[2,-3,2,0,-4],[-3,3,0,3,4]),
         ([-5,6,-2,3,8],[5,-3,2,0,-4],[3,0,3,4]),
         ([-5,6,-2,3,8],[5,-6,2,0,-4],[3,4])]
for (p1, p2, resultado) in tests:
    print('   polinomio1: {} polinomio2: {} resultado: {}'.format(p1, p2, resultado))
    p1_copia=p1.copy()
    p2_copia=p2.copy()
    result=polynomial_sum(p1_copia, p2_copia)
    try:
        assert resultado==result and p1==p1_copia and p2==p2_copia
    except:
        print('      RESULTADO INCORRECTO / MODIFICACIÓN DE PARÁMETROS - polinomio1: {} polinomio2: {} resultado: {}'.format(p1_copia, p2_copia, result))
print('FIN de las pruebas (polynomial_sum())')
### END TESTS
```

Output:
```
INICIO de las pruebas (polynomial_sum())
   polinomio1: [] polinomio2: [3] resultado: [3]
   polinomio1: [-8] polinomio2: [-5, 6, -2, 3, 8] resultado: [-5, 6, -2, 3, 0]
   polinomio1: [1, 1] polinomio2: [] resultado: [1, 1]
   polinomio1: [1, 0, -1] polinomio2: [-1, 0, 1] resultado: []
   polinomio1: [1, 5, 0, 2] polinomio2: [-3, 0] resultado: [1, 5, -3, 2]
   polinomio1: [1, 2, 1, 1] polinomio2: [3, 2, 1] resultado: [1, 5, 3, 2]
   polinomio1: [-5, 6, -2, 3, 8] polinomio2: [2, -3, 2, 0, -4] resultado: [-3, 3, 0, 3, 4]
   polinomio1: [-5, 6, -2, 3, 8] polinomio2: [5, -3, 2, 0, -4] resultado: [3, 0, 3, 4]
   polinomio1: [-5, 6, -2, 3, 8] polinomio2: [5, -6, 2, 0, -4] resultado: [3, 4]
FIN de las pruebas (polynomial_sum())
```

### Ejercicio 3
Diseña una función llamada `polynomial_product(p1,p2)` que, dados dos polinomios `p1` y `p2`, calcule el producto `p1*p2`. Por ejemplo, el resultado de la llamada `polynomial_product([1,2,1,1],[3,2,1])` debe ser la lista de coeficientes correspondiente al polinomio $q(x)\!=\!3x^5\!+\!8x^4\!+\!8x^3\!+\!7x^2\!+\!3x\!+\!1$.

<!-- grade_id: product_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def polynomial_product(p1,p2):
    gradoResultado=(len(p1)-1)+(len(p2)-1)
    producto=[0]*(gradoResultado+1)
    for i in range(len(p1)):
        for j in range(len(p2)):
            producto[i+j]=producto[i+j]+p1[i]*p2[j]
    polynomial_normalization(producto)
    return producto
### END SOLUTION
```

<!-- grade_id: product_test, points: 3, grade: true, solution: false -->
```python
### BEGIN TESTS
print('INICIO de las pruebas (polynomial_product())')
tests = [([],[3,-1],[]),([2,-1],[2,-1],[4,-4,1]),
         ([1,1,1],[],[]),
         ([1,1,-2],[3,7],[3,10,1,-14]),
         ([-1,0,2],[3,1],[-3,-1,6,2]),
         ([1,0,-1],[1,0,1],[1,0,0,0,-1]),
         ([3,-7,-1],[2],[6,-14,-2]),
         ([3,-7,-1],[-4,0,0,0,0,0],[-12,28,4,0,0,0,0,0]),
         ([-4,0,0,0,0,0],[3,-7,-1],[-12,28,4,0,0,0,0,0])]
for (p1, p2, resultado) in tests:
    print('   polinomio1: {} polinomio2: {} resultado: {}'.format(p1, p2, resultado))
    p1_copia=p1.copy()
    p2_copia=p2.copy()
    result=polynomial_product(p1_copia, p2_copia)
    try:
        assert resultado==result and p1==p1_copia and p2==p2_copia
    except:
        print('      RESULTADO INCORRECTO / MODIFICACIÓN DE PARÁMETROS - polinomio1: {} polinomio2: {} resultado: {}'.format(p1_copia, p2_copia, result))
print('FIN de las pruebas (polynomial_product())')
### END TESTS
```

Output:
```
INICIO de las pruebas (polynomial_product())
   polinomio1: [] polinomio2: [3, -1] resultado: []
   polinomio1: [2, -1] polinomio2: [2, -1] resultado: [4, -4, 1]
   polinomio1: [1, 1, 1] polinomio2: [] resultado: []
   polinomio1: [1, 1, -2] polinomio2: [3, 7] resultado: [3, 10, 1, -14]
   polinomio1: [-1, 0, 2] polinomio2: [3, 1] resultado: [-3, -1, 6, 2]
   polinomio1: [1, 0, -1] polinomio2: [1, 0, 1] resultado: [1, 0, 0, 0, -1]
   polinomio1: [3, -7, -1] polinomio2: [2] resultado: [6, -14, -2]
   polinomio1: [3, -7, -1] polinomio2: [-4, 0, 0, 0, 0, 0] resultado: [-12, 28, 4, 0, 0, 0, 0, 0]
   polinomio1: [-4, 0, 0, 0, 0, 0] polinomio2: [3, -7, -1] resultado: [-12, 28, 4, 0, 0, 0, 0, 0]
FIN de las pruebas (polynomial_product())
```

### Ejercicio 4
Diseña una función llamada `polynomial_ruffiniRule(p,n)` que, dados un polinomio `p` y un número entero `n`, devuelva el cociente y el resto de la división de `p` entre el polinomio `x-n` aplicando la regla de Ruffini. Por ejemplo, el resultado de la llamada `polynomial_ruffiniRule([6,-3,0,2,-1],-1)` debe ser la tupla compuesta por la lista de coeficientes correspondiente al polinomio $q(x)\!=\!6x^3\!-\!9x^2\!+\!9x\!-\!7$ y el resto $6$, ya que $6x^4\!-\!3x^3\!+\!2x-1\!=\!(6x^3\!-\!9x^2\!+\!9x\!-\!7)(x\!+\!1)\!+\!6$.  
**Observación:** Debe comprobarse que el parámetro `n` es un divisor del término independiente de `p`.

<!-- grade_id: ruffiniRule_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def polynomial_ruffiniRule(dividendo,divisor):
    if (len(dividendo)==0) or (dividendo[-1]%divisor!=0):
        raise ValueError("Parametro incorrecto")
    cociente=[dividendo[0]]
    for paso in range(1,len(dividendo)):
        cociente.append(cociente[-1]*divisor+dividendo[paso])
    resto=cociente.pop(-1)
    return cociente,resto
### END SOLUTION
```

<!-- grade_id: ruffiniRule_test, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
print('INICIO de las pruebas (polynomial_ruffiniRule())')
tests = [([-17],1,([],-17)),([1,1],-1,([1],0)),([1,3,2],-2,([1,1],0)),
         ([1,2,-1,-2],1,([1,3,2],0)),
         ([1,2,-1,-2],-2,([1,0,-1],0)),
         ([6,0,0,0],1,([6,6,6],6)),
         ([1,1,-6,-4,8],1,([1,2,-4,-8],0)),
         ([3,0,0,0,2,4],-2,([3,-6,12,-24,50],-96)),
         ([6,-3,0,0,2,0],-1,([6,-9,9,-9,11],-11)),
         ([3,-4,0,-6,-7,0],-2,([3,-10,20,-46,85],-170))]
for (p, n, resultado) in tests:
    print('   polinomio: {} valor: {} resultado: {}'.format(p, n, resultado))
    p_copia=p.copy()
    result=polynomial_ruffiniRule(p_copia, n)
    try:
        assert resultado==result and p==p_copia
    except:
        print('      RESULTADO INCORRECTO / MODIFICACIÓN DE PARÁMETROS - polinomio: {} resultado: {}'.format(p_copia, result))
print('FIN de las pruebas (polynomial_ruffiniRule())')
### END TESTS
```

Output:
```
INICIO de las pruebas (polynomial_ruffiniRule())
   polinomio: [-17] valor: 1 resultado: ([], -17)
   polinomio: [1, 1] valor: -1 resultado: ([1], 0)
   polinomio: [1, 3, 2] valor: -2 resultado: ([1, 1], 0)
   polinomio: [1, 2, -1, -2] valor: 1 resultado: ([1, 3, 2], 0)
   polinomio: [1, 2, -1, -2] valor: -2 resultado: ([1, 0, -1], 0)
   polinomio: [6, 0, 0, 0] valor: 1 resultado: ([6, 6, 6], 6)
   polinomio: [1, 1, -6, -4, 8] valor: 1 resultado: ([1, 2, -4, -8], 0)
   polinomio: [3, 0, 0, 0, 2, 4] valor: -2 resultado: ([3, -6, 12, -24, 50], -96)
   polinomio: [6, -3, 0, 0, 2, 0] valor: -1 resultado: ([6, -9, 9, -9, 11], -11)
   polinomio: [3, -4, 0, -6, -7, 0] valor: -2 resultado: ([3, -10, 20, -46, 85], -170)
FIN de las pruebas (polynomial_ruffiniRule())
```

### Ejercicio 5
Diseña una función llamada `polynomial_bulkSum(lst)` que, dada una lista de polinomios (no necesariamente del mismo grado), calcule la suma de los mismos. Por ejemplo, el resultado de la llamada `polynomial_bulkSum([[1,2,1,1],[3,2,1],[-1,0,0,0]])` debe ser la lista de coeficientes correspondiente al polinomio $q(x)\!=\!5x^2\!+\!3x\!+\!2$.

<!-- grade_id: bulkSum_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def polynomial_bulkSum(lst):
    suma=[]
    for polinomio in lst:
        suma=polynomial_sum(suma,polinomio)
    return suma
### END SOLUTION
```

<!-- grade_id: bulkSum_test, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
print('INICIO de las pruebas (polynomial_bulkSum())')
tests = [([],[]),([[]],[]),([[1,-9]],[1,-9]),
         ([[2,5,-6],[3,-5,3],[3],[]],[5,0,0]),
         ([[2,5,-6],[3,-5,3],[3],[],[-5,0,0]],[]),
         ([[3,-7,-1],[-4,0,0,0,0,0]],[-4,0,0,3,-7,-1]),
         ([[3,1,-2,1,-14],[6,-8,2,-3,0],[2,14]],[9,-7,0,0,0]),
         ([[6,-3,8,4,5],[-9,5,6,-2,7],[-4,6,-9,6,-3]],[-7,8,5,8,9]),
         ([[-1,4,-2,-7,6,1,-2],[3,2,0,-1,2,-2,5],[-2,-6,2,8,-8,1,-3]],[]),
         ([[-1,4,-2,-7,6,1,-2],[3,2,0,-1,2,-2,5],[-2,-6,2,8,-8,1,-3],[5,5,5]],[5,5,5])]
for (lst, resultado) in tests:
    print('   polinomios: {} resultado: {}'.format(lst, resultado))
    lst_copia=[] # copia profunda
    for polinomio in lst:
        lst_copia.append(polinomio.copy())
    result=polynomial_bulkSum(lst_copia)
    try:
        assert resultado==result and lst==lst_copia
    except:
        print('      RESULTADO INCORRECTO / MODIFICACIÓN DE PARÁMETROS - polinomios: {} resultado: {}'.format(lst_copia, result))
print('FIN de las pruebas (polynomial_bulkSum())')
### END TESTS
```

Output:
```
INICIO de las pruebas (polynomial_bulkSum())
   polinomios: [] resultado: []
   polinomios: [[]] resultado: []
   polinomios: [[1, -9]] resultado: [1, -9]
   polinomios: [[2, 5, -6], [3, -5, 3], [3], []] resultado: [5, 0, 0]
   polinomios: [[2, 5, -6], [3, -5, 3], [3], [], [-5, 0, 0]] resultado: []
   polinomios: [[3, -7, -1], [-4, 0, 0, 0, 0, 0]] resultado: [-4, 0, 0, 3, -7, -1]
   polinomios: [[3, 1, -2, 1, -14], [6, -8, 2, -3, 0], [2, 14]] resultado: [9, -7, 0, 0, 0]
   polinomios: [[6, -3, 8, 4, 5], [-9, 5, 6, -2, 7], [-4, 6, -9, 6, -3]] resultado: [-7, 8, 5, 8, 9]
   polinomios: [[-1, 4, -2, -7, 6, 1, -2], [3, 2, 0, -1, 2, -2, 5], [-2, -6, 2, 8, -8, 1, -3]] resultado: []
   polinomios: [[-1, 4, -2, -7, 6, 1, -2], [3, 2, 0, -1, 2, -2, 5], [-2, -6, 2, 8, -8, 1, -3], [5, 5, 5]] resultado: [5, 5, 5]
FIN de las pruebas (polynomial_bulkSum())
```