## Círculos y circunferencias

Realiza una función que calcule la longiutd de una circunferencia de un radio dado

```python
import math
```

<!-- grade_id: cell-f1da7b2c38093314, points: 0, grade: false, solution: true -->
```python

def circumference(radius: float) -> float:
    ### BEGIN SOLUTION
    return 2 * radius * math.pi
    ### END SOLUTION
```

```python
### Escribe aquí tus pruebas

```

<!-- grade_id: cell-27dea47fbdfe02ba, points: 1, grade: true, solution: false -->
```python
def test_circumference():
    tests = [(0.0, 0.0),
            (0.5, 3.141592653589793),
            (1.0, 6.283185307179586),
            (1.5, 9.42477796076938),
            (2.0, 12.566370614359172),
            (2.5, 15.707963267948966),
            (3.0, 18.84955592153876),
            (3.5, 21.991148575128552),
            (4.0, 25.132741228718345),
            (4.5, 28.274333882308138),
            (5.0, 31.41592653589793)]
    for rad, length in tests:
        print(f'Testing {rad} {length}....', end='')
        res = circumference(rad)
        assert math.isclose(res, length), f'Error {rad} {length} {res}'
        print('OK')

test_circumference()
```

Output:
```
Testing 0.0 0.0....OK
Testing 0.5 3.141592653589793....OK
Testing 1.0 6.283185307179586....OK
Testing 1.5 9.42477796076938....OK
Testing 2.0 12.566370614359172....OK
Testing 2.5 15.707963267948966....OK
Testing 3.0 18.84955592153876....OK
Testing 3.5 21.991148575128552....OK
Testing 4.0 25.132741228718345....OK
Testing 4.5 28.274333882308138....OK
Testing 5.0 31.41592653589793....OK
```

Escribe ahora una función que calcule la superficie de un círculo de un radio dado

<!-- grade_id: cell-6ea9edcad255937b, points: 0, grade: false, solution: true -->
```python
def circle(radius: float) -> float:
    ### BEGIN SOLUTION
    return math.pi * radius**2
    ### END SOLUTION

```

```python
### Escribe aquí tus pruebas

```

<!-- grade_id: cell-6875537885715192, points: 1, grade: true, solution: false -->
```python
def test_circle():
    tests = [(0.0, 0.0),
             (0.5, 0.7853981633974483),
             (1.0, 3.141592653589793),
             (1.5, 7.0685834705770345),
             (2.0, 12.566370614359172),
             (2.5, 19.634954084936208),
             (3.0, 28.274333882308138),
             (3.5, 38.48451000647496),
             (4.0, 50.26548245743669),
             (4.5, 63.61725123519331),
             (5.0, 78.53981633974483)]
    for rad, surface in tests:
        print(f'Testing {rad} {surface}....', end='')
        res = circle(rad)
        assert math.isclose(res, surface), f'Error {rad} {surface} {res}'
        print('OK')

def test_circle()
```