# Término general de una sucesión.

Encuentra una expresión en python que sean el término general de las siguientes sucesiones. 

0, 1, 0, 1, 0, 1, ....

<!-- grade_id: cell-24b48743b02321e3, points: 0, grade: false, solution: true -->
```python
def zero_one(n: int) -> int:
    ### BEGIN SOLUTION
    return n % 2
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-5c181fd49e5f56ae, points: 1, grade: true, solution: false -->
```python
def test_zero_one():
    result = list(map(zero_one, range(6)))
    assert result == [0, 1, 0, 1, 0, 1], f'Error {result} is incorrect'
    print('OK')             
test_zero_one()
```

Output:
```
OK
```

1, -1, 1, -1, 1, -1, ....

Intenta encontrar una expresión una expresión que no involucre una potencia: $-1^n$. Usa el operador `%`.

<!-- grade_id: cell-c63ef72c9ecbb0b0, points: 0, grade: false, solution: true -->
```python
def signo(n: int) -> int:
    ### BEGIN SOLUTION
    return 1 - 2 * (n % 2)
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
```

```python
def test_signo():
    result = list(map(signo, range(6)))
    assert result == [1, -1, 1, -1, 1, -1], f'Error {result} is incorrect'
    print('OK')             
test_signo()
```

Output:
```
OK
```

-1, 1, -1, 1, -1, 1,

<!-- grade_id: cell-2a8159ac238d391c, points: 0, grade: false, solution: true -->
```python
def signo2(n: int) -> int:
    ### BEGIN SOLUTION
    return -1 + 2 * (n % 2)
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-54dcce97dbf074be, points: 1, grade: true, solution: false -->
```python
def test_signo2():
    result = list(map(signo2, range(6)))
    assert result == [-1, 1, -1, 1, -1, 1], f'Error {result} is incorrect'
    print('OK')             
test_signo2()
```

Output:
```
OK
```

 3, 8, 13, 18, 23, .....

<!-- grade_id: cell-feb2cade1fa63009, points: 0, grade: false, solution: true -->
```python
def lineal(n: int) -> int:
    ### BEGIN SOLUTION
    return 3 + 5*n
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-ea572850101eddf7, points: 1, grade: true, solution: false -->
```python
def test_lineal():
    result = list(map(lineal, range(6)))
    assert result == [3, 8, 13, 18, 23, 28], f'Error {result} is incorrect'
    print('OK')             
test_lineal()
```

Output:
```
OK
```

0, 1, 2, 3, 0, 1, 2, 3, 0, 1, 2, 3, ....

<!-- grade_id: cell-53df024e93817f97, points: 0, grade: false, solution: true -->
```python
def modulo(n: int) -> int:
    ### BEGIN SOLUTION
    return n%4
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-479eebd9215d1f3f, points: 1, grade: true, solution: false -->
```python
def test_modulo():
    result = list(map(modulo, range(12)))
    assert result == [0, 1, 2, 3, 0, 1, 2, 3, 0, 1, 2, 3], f'Error {result} is incorrect'
    print('OK')             
test_modulo()
```

Output:
```
OK
```

1, 1, 2, 3, 5, 8, 13, 21 ...

Consulta un término general para la sucesión de Fibonacci 

<img src='Fibonacci_spiral_34.svg' width='250px'> 
<img src='nautilius.jpg' width='250px'> 

[https://es.wikipedia.org/wiki/Sucesi%C3%B3n_de_Fibonacci]: 

<!-- grade_id: cell-3aed57907e9717d9, points: 0, grade: false, solution: true -->
```python
def fibonacci(n: int) -> int:
    ### BEGIN SOLUTION
    import math
    phi = (1 + math.sqrt(5)) / 2
    return int(phi**(n+1)/math.sqrt(5) + 0.5)
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas

```

Output:
```
2
```

```python
def test_fibonacci():
    result = list(map(fibonacci, range(8)))
    assert result == [1, 1, 2, 3, 5, 8, 13, 21], f'Error {result} is incorrect'
    print('OK')             
test_fibonacci()
```

Output:
```
OK
```