# Conversiones de unidades

Como todos sabemos el mundo anglosajón es reacio a cambiar sus viejas costumbres. Por ello, siglos después de que todo el mundo haya optado por usar el sistema métrico decimal, ellos insisten en usar su arcaico sistema de unidades. 

## Distancias.
Vamos a empezar con las unidades de distancia.

* 1 pulgada (in) = 2.54 cm
* 1 pie (ft) = 12 in = 30.48 cm
* 1 yarda (yd) = 3 ft = 91.44 cm
* 1 milla (mi) = 1760 yd = 1.609344 km

En este ejercico vamos a realizar funciones que nos ayuden a hacer entender su sistema de unidades.

Escribe una función que convierta pulgadas a centimetros `in2cm` y otra que convierta de centímetros a pulgadas `cm2in`. Análogamente:
* de pies a metros `ft2m` y `m2ft`
* de yardas a metros `yd2m` y `m2yd`
* de millas a kilómetros `mi2km` y `km2mi`


<!-- grade_id: cell-8f50da666022e0ff, points: 0, grade: false, solution: true -->
```python
def in2cm(inches: float) -> float:
    ### BEGIN SOLUTION
    return inches * 2.54
    ### END SOLUTION
    
def cm2in(centimeters: float) -> float:
    ### BEGIN SOLUTION
    return centimeters / 2.54
    ### END SOLUTION
    

```

<!-- grade_id: cell-75c3aeafe04d90d5, points: 0, grade: false, solution: true -->
```python
### Escribe aquí tus pruebas
```

<!-- grade_id: cell-e057df4058573bc6, points: 1, grade: true, solution: false -->
```python
def test_inches():
    import math
    print('Testing inches............')
    tests = [(-5, -12.7), 
             (-4, -10.16),
             (-3, -7.62),
             (-2, -5.08),
             (-1, -2.54),
             (0, 0.0),
             (1, 2.54),
             (2, 5.08),
             (3, 7.62),
             (4, 10.16),
             (5, 12.7),
             (-1.968503937007874, -5),
             (-1.574803149606299, -4),
             (-1.1811023622047243, -3),
             (-0.7874015748031495, -2),
             (-0.39370078740157477, -1),
             (0.0, 0),
             (0.39370078740157477, 1),
             (0.7874015748031495, 2),
             (1.1811023622047243, 3),
             (1.574803149606299, 4),
             (1.968503937007874, 5)]
    for inches, centimeters in tests:
        print(f'in:{inches:+8.4f} cm:{centimeters: 8.4f}....', end='')
        cm = in2cm(inches)
        assert math.isclose(cm, centimeters), f'Error in2cm({inches})={cm}'
        inches_r = cm2in(centimeters)
        assert math.isclose(inches, inches_r), f'Error mc2in({centimeters})={inches_r}'
        print('OK')
        
test_inches()
            
        
```

Output:
```
Testing inches............
in: -5.0000 cm:-12.7000....OK
in: -4.0000 cm:-10.1600....OK
in: -3.0000 cm: -7.6200....OK
in: -2.0000 cm: -5.0800....OK
in: -1.0000 cm: -2.5400....OK
in: +0.0000 cm:  0.0000....OK
in: +1.0000 cm:  2.5400....OK
in: +2.0000 cm:  5.0800....OK
in: +3.0000 cm:  7.6200....OK
in: +4.0000 cm: 10.1600....OK
in: +5.0000 cm: 12.7000....OK
in: -1.9685 cm: -5.0000....OK
in: -1.5748 cm: -4.0000....OK
in: -1.1811 cm: -3.0000....OK
in: -0.7874 cm: -2.0000....OK
in: -0.3937 cm: -1.0000....OK
in: +0.0000 cm:  0.0000....OK
in: +0.3937 cm:  1.0000....OK
in: +0.7874 cm:  2.0000....OK
in: +1.1811 cm:  3.0000....OK
in: +1.5748 cm:  4.0000....OK
in: +1.9685 cm:  5.0000....OK
```

<!-- grade_id: cell-cc989805b5fac988, points: 0, grade: false, solution: true -->
```python
def ft2m(feet: float) -> float:
    ### BEGIN SOLUTION
    return feet * 0.3048
    ### END SOLUTION

def m2ft(meters: float) -> float:
    ### BEGIN SOLUTION
    return meters / 0.3048
    ### END SOLUTION


```

```python
#Escribe aquí tus pruebas
```

<!-- grade_id: cell-20fb724e090610ec, points: 1, grade: true, solution: false -->
```python
def test_feet():
    import math
    print('Testing feet............')
    tests = [(-5, -1.524),
             (-4, -1.2192),
             (-3, -0.9144000000000001),
             (-2, -0.6096),
             (-1, -0.3048),
             (0, 0.0),
             (1, 0.3048),
             (2, 0.6096),
             (3, 0.9144000000000001),
             (4, 1.2192),
             (5, 1.524),
             (-16.404199475065617, -5),
             (-13.123359580052492, -4),
             (-9.84251968503937, -3),
             (-6.561679790026246, -2),
             (-3.280839895013123, -1),
             (0.0, 0),
             (3.280839895013123, 1),
             (6.561679790026246, 2),
             (9.84251968503937, 3),
             (13.123359580052492, 4),
             (16.404199475065617, 5)]

    for feet, centimeters in tests:
        print(f'feet:{feet:+8.4f} meters:{centimeters: 8.4f}....', end='')
        meters = ft2m(feet)
        assert math.isclose(meters, centimeters), f'Error ft2m({feet})={meters}'
        feet_r = m2ft(centimeters)
        assert math.isclose(feet, feet_r), f'Error m2ft({centimeters})={feet_r}'
        print('OK')

test_feet()

```

Output:
```
Testing feet............
feet: -5.0000 meters: -1.5240....OK
feet: -4.0000 meters: -1.2192....OK
feet: -3.0000 meters: -0.9144....OK
feet: -2.0000 meters: -0.6096....OK
feet: -1.0000 meters: -0.3048....OK
feet: +0.0000 meters:  0.0000....OK
feet: +1.0000 meters:  0.3048....OK
feet: +2.0000 meters:  0.6096....OK
feet: +3.0000 meters:  0.9144....OK
feet: +4.0000 meters:  1.2192....OK
feet: +5.0000 meters:  1.5240....OK
feet:-16.4042 meters: -5.0000....OK
feet:-13.1234 meters: -4.0000....OK
feet: -9.8425 meters: -3.0000....OK
feet: -6.5617 meters: -2.0000....OK
feet: -3.2808 meters: -1.0000....OK
feet: +0.0000 meters:  0.0000....OK
feet: +3.2808 meters:  1.0000....OK
feet: +6.5617 meters:  2.0000....OK
feet: +9.8425 meters:  3.0000....OK
feet:+13.1234 meters:  4.0000....OK
feet:+16.4042 meters:  5.0000....OK
```

<!-- grade_id: cell-f4e17e8ac95e4a22, points: 0, grade: false, solution: true -->
```python
def yd2m(yards: float) -> float:
    ### BEGIN SOLUTION
    return yards * 0.9144
    ### END SOLUTION

def m2yd(meters: float) -> float:
    ### BEGIN SOLUTION
    return meters / 0.9144
    ### END SOLUTION


```

```python
### Escribe aquí tus pruebas

```

```python
def test_yards():
    import math
    print('Testing yards............')
    tests = [(-5, -4.572),
             (-4, -3.6576),
             (-3, -2.7432),
             (-2, -1.8288),
             (-1, -0.9144),
             (0, 0.0),
             (1, 0.9144),
             (2, 1.8288),
             (3, 2.7432),
             (4, 3.6576),
             (5, 4.572),
             (-5.468066491688539, -5),
             (-4.374453193350831, -4),
             (-3.2808398950131235, -3),
             (-2.1872265966754156, -2),
             (-1.0936132983377078, -1),
             (0.0, 0),
             (1.0936132983377078, 1),
             (2.1872265966754156, 2),
             (3.2808398950131235, 3),
             (4.374453193350831, 4),
             (5.468066491688539, 5)]

    for yards, centimeters in tests:
        print(f'yards:{yards:+8.4f} meters:{centimeters: 8.4f}....', end='')
        meters = yd2m(yards)
        assert math.isclose(meters, centimeters), f'Error yd2m({yards})={meters}'
        yards_r = m2yd(centimeters)
        assert math.isclose(yards, yards_r), f'Error m2yd({centimeters})={yards_r}'
        print('OK')

test_yards()

```

Output:
```
Testing yards............
yards: -5.0000 meters: -4.5720....OK
yards: -4.0000 meters: -3.6576....OK
yards: -3.0000 meters: -2.7432....OK
yards: -2.0000 meters: -1.8288....OK
yards: -1.0000 meters: -0.9144....OK
yards: +0.0000 meters:  0.0000....OK
yards: +1.0000 meters:  0.9144....OK
yards: +2.0000 meters:  1.8288....OK
yards: +3.0000 meters:  2.7432....OK
yards: +4.0000 meters:  3.6576....OK
yards: +5.0000 meters:  4.5720....OK
yards: -5.4681 meters: -5.0000....OK
yards: -4.3745 meters: -4.0000....OK
yards: -3.2808 meters: -3.0000....OK
yards: -2.1872 meters: -2.0000....OK
yards: -1.0936 meters: -1.0000....OK
yards: +0.0000 meters:  0.0000....OK
yards: +1.0936 meters:  1.0000....OK
yards: +2.1872 meters:  2.0000....OK
yards: +3.2808 meters:  3.0000....OK
yards: +4.3745 meters:  4.0000....OK
yards: +5.4681 meters:  5.0000....OK
```

<!-- grade_id: cell-08030cecc7c9cc73, points: 0, grade: false, solution: true -->
```python
def mi2km(miles: float) -> float:
    ### BEGIN SOLUTION
    return miles * 1.609344
    ### END SOLUTION

def km2mi(kilometers: float) -> float:
    ### BEGIN SOLUTION
    return kilometers / 1.609344
    ### END SOLUTION

```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-ab61d23dd739eb24, points: 1, grade: true, solution: false -->
```python
def test_miles():
    import math
    print('Testing miles............')
    tests = [(-5, -8.04672),
             (-4, -6.437376),
             (-3, -4.828032),
             (-2, -3.218688),
             (-1, -1.609344),
             (0, 0.0),
             (1, 1.609344),
             (2, 3.218688),
             (3, 4.828032),
             (4, 6.437376),
             (5, 8.04672),
             (-3.1068559611866697, -5),
             (-2.485484768949336, -4),
             (-1.8641135767120018, -3),
             (-1.242742384474668, -2),
             (-0.621371192237334, -1),
             (0.0, 0),
             (0.621371192237334, 1),
             (1.242742384474668, 2),
             (1.8641135767120018, 3),
             (2.485484768949336, 4),
             (3.1068559611866697, 5)]

    for miles, kilometers in tests:
        print(f'miles:{miles:+8.4f} meters:{kilometers: 8.4f}....', end='')
        meters = mi2km(miles)
        assert math.isclose(meters, kilometers), f'Error mi2km({miles})={meters}'
        miles_r = km2mi(kilometers)
        assert math.isclose(miles, miles_r), f'Error km2yd({kilometers})={miles_r}'
        print('OK')

test_miles()

```

Output:
```
Testing miles............
miles: -5.0000 meters: -8.0467....OK
miles: -4.0000 meters: -6.4374....OK
miles: -3.0000 meters: -4.8280....OK
miles: -2.0000 meters: -3.2187....OK
miles: -1.0000 meters: -1.6093....OK
miles: +0.0000 meters:  0.0000....OK
miles: +1.0000 meters:  1.6093....OK
miles: +2.0000 meters:  3.2187....OK
miles: +3.0000 meters:  4.8280....OK
miles: +4.0000 meters:  6.4374....OK
miles: +5.0000 meters:  8.0467....OK
miles: -3.1069 meters: -5.0000....OK
miles: -2.4855 meters: -4.0000....OK
miles: -1.8641 meters: -3.0000....OK
miles: -1.2427 meters: -2.0000....OK
miles: -0.6214 meters: -1.0000....OK
miles: +0.0000 meters:  0.0000....OK
miles: +0.6214 meters:  1.0000....OK
miles: +1.2427 meters:  2.0000....OK
miles: +1.8641 meters:  3.0000....OK
miles: +2.4855 meters:  4.0000....OK
miles: +3.1069 meters:  5.0000....OK
```

## Temperaturas
Vamos a seguir con las unidades de temperatura. El grado Fahrenheit (representado como °F) es una escala de temperatura propuesta por el físico e ingeniero polaco Daniel Gabriel Fahrenheit en 1724. La escala establece como las temperaturas de congelación y ebullición del agua, 32 °F y 212 °F, respectivamente. El método de definición es similar al utilizado para el grado Celsius (°C) [https://es.wikipedia.org/wiki/Grado_Fahrenheit].

Realiza una función `fhr2cel` que transforme los gradeos Fahrenhiet a Celsius, y la inversa `cel2fhr` que realize la transofrmación inversa. 

<!-- grade_id: cell-5fbaf6dfdae387cb, points: 0, grade: false, solution: true -->
```python
def fhr2cel(fhr: float) -> float:
    ### BEGIN SOLUTION
    return ( fhr-32 )/1.8
    ### END SOLUTION

def cel2fhr(cel: float) -> float:
    ### BEGIN SOLUTION
    return 32 +  cel * 1.8
    ### END SOLUTION

```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-ee23e7c673f5833a, points: 1, grade: true, solution: false -->
```python
def test_fhr():
    import math
    print('Testing fhr............')
    tests = [(-40, -40.0),
             (-30, -34.44444444444444),
             (-20, -28.88888888888889),
             (-10, -23.333333333333332),
             (0, -17.77777777777778),
             (10, -12.222222222222221),
             (20, -6.666666666666666),
             (30, -1.1111111111111112),
             (40, 4.444444444444445),
             (50, 10.0),
             (60, 15.555555555555555),
             (70, 21.11111111111111),
             (80, 26.666666666666664),
             (90, 32.22222222222222),
             (-40.0, -40),
             (-22.0, -30),
             (-4.0, -20),
             (14.0, -10),
             (32.0, 0),
             (50.0, 10),
             (68.0, 20),
             (86.0, 30),
             (104.0, 40),
             (122.0, 50),
             (140.0, 60),
             (158.0, 70),
             (176.0, 80),
             (194.0, 90)]

    for fhr, cel in tests:
        print(f'fhr:{fhr:+8.4f} cel:{cel: 8.4f}....', end='')
        cel = fhr2cel(fhr)
        assert math.isclose(cel, cel), f'Error mi2km({fhr})={cel}'
        fhr_r = cel2fhr(cel)
        assert math.isclose(fhr, fhr_r), f'Error km2yd({cel})={fhr_r}'
        print('OK')

```