# Expresiones de triángulos

En este ejercicio vamos a realizar funciones que comprueben si una terna de magnitudes (positivas) pueden formar un triángulo.

<img src='./triangulo.png'>

Para que tres magnitudes positivas puedan formar un triangulo es necesario y suficiente que la suma de dos de ellas sea superior a la tercera.
Escribe una función llamada `is_triangle` que compruebe si tres magnitudes pueden ser las medidas de los lados de un triangulos.

<!-- grade_id: cell-707768d95e73b433, points: 0, grade: false, solution: true -->
```python
def is_triangle(a: float, b: float, c: float) -> bool:
    ### BEGIN SOLUTION
    return a>0 and b>0  and c>0 and (a+b)>c and (a+c)>b and (c+b)>a
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
is_triangle(3,4,5), is_triangle(-1,-1,-1), is_triangle(2,3,10)
```

<!-- grade_id: cell-192ef82fca77e11c, points: 1, grade: true, solution: false -->
```python
def runtest(tests, fun):
    for i, o in tests:
        print(f'{i}, {o}....', end='')
        res = fun(*i)
        assert  res == o, f'error input:{i}, out: {res}, expected: {o}' 
        print('OK')
    print('OK')
    
def test_is_triangle():
    tests = [
        ((3,4,5), True),
        ((1,2,4), False),
        ((4,1,2), False),
        ((1,4,2), False),
        ((-1,2,3), False),
        ((1,-2,3), False),
    ]
    runtest(tests, is_triangle)
        
test_is_triangle()
```

## Triángulo equilátero

Realiza una función llamada `is_triangle` que compruebe si tres magnitudes pueden ser los lados de un triángulo equilátero

<!-- grade_id: cell-b6e9e403b20da31c, points: 0, grade: false, solution: true -->
```python
def is_equilateral(a: float, b: float, c: float) -> bool:
    ### BEGIN SOLUTION
    import math
    return is_triangle(a, b, c) and math.isclose(a, b) and math.isclose(b, c)
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
is_equilateral(1,1,1), is_equilateral(-1, -1, -1), is_equilateral(4,2,2)
```

<!-- grade_id: cell-a3cd6f8ac645db55, points: 1, grade: true, solution: false -->
```python
def test_is_equilateral():
    tests = [
        ((1,1,1), True),
        ((2,2,3), False),
        ((3,2,2), False),
        ((2,3,2), False),
        ((1,2,4), False),
        ((-1, -1, -1), False)
    ]
    test_fun(tests, is_equilateral)
        
test_is_equilateral()
```

## Triángulo Isósceles

Haz ahora una función llamada `is_isosceles` que compruebe si tres magnitudes dadas pueden formar una triángulo isósceles.

<!-- grade_id: cell-af62c0c1394a5e05, points: 0, grade: false, solution: true -->
```python
def is_isosceles(a: float, b: float, c: float) -> bool:
    ### BEGIN SOLUTION
    import math
    return is_triangle(a, b, c) and \
        (math.isclose(a, b) or math.isclose(a, c) or math.isclose(b, c))
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
is_isosceles(1,1,1), is_isosceles(-1, -1, -1), is_isosceles(4,2,2)
```

<!-- grade_id: cell-14357f1df455427b, points: 1, grade: true, solution: false -->
```python
def test_is_isosceles():
    tests = [
        ((1,1,1), True),
        ((2,2,3), True),
        ((3,2,2), True),
        ((2,3,2), True),
        ((3,4,5), False),
        ((1,1,3), False),
        ((-1,-1, 1), False),
    ]
    test_fun(tests, is_isosceles)
        
test_is_isosceles()
```

## Triángulo Escaleno

Haz ahora una función llamada `is_scalene` que compruebe si tres magnitudes dadas pueden formar una triángulo escaleno.

<!-- grade_id: cell-c52187c21aa2805e, points: 0, grade: false, solution: true -->
```python
def is_scalene(a: float, b: float, c: float) -> bool:
    ### BEGIN SOLUTION
        import math
        return is_triangle(a, b, c) and \
        not math.isclose(a, b) and\
        not math.isclose(a, c) and\
        not math.isclose(b, c)
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
is_scalene(3,4,6), is_scalene(3,4,10), is_scalene(4,2,2)
```

<!-- grade_id: cell-ae4eb05977d3e67f, points: 1, grade: true, solution: false -->
```python
def test_is_scalene():
    tests = [
        ((3,4,5), True),
        ((-3,4,5), False),
        ((1,1,1), False),
        ((2,2,3), False),
        ((3,2,2), False),
        ((2,3,2), False),
        ((1,1,3), False),
        ((1,2,5), False)
    ]
    test_fun(tests, is_scalene)
        
test_is_scalene()
```