# DNI

En este ejercicio vamos a calcular la letra correspondiente a un Documento Nacional de Ientidad (DNI) l

<img src='dni.jpg'>

Para calcular la letra correspondiente al un número, basta dividir este por 23 y fijarse en el resto. La letra se averigua consultando la siguiente tabla

| 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | 9 |10 |11 |12 |13 |14 |15 |16 |17 |18 |19 |10 |11 |12 | 
|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|---|
| T | R | W | A | G | M | Y | F | P | D | X | B | N | J | Z | S | Q | V | H | L | C | K | E |

Realiza una función llamada `letra` que dado un número de DNI, nos calcule la letra correspondiente. 
Pista, conviene usar la cadena de caracteres `TRWAGMYFPDXBNJZSQVHLCKE`


<!-- grade_id: cell-c7f7eb71c3270c46, points: 0, grade: false, solution: true -->
```python
def letra(dni: int) -> str:
    ### BEGIN SOLUTION
    tabla = "TRWAGMYFPDXBNJZSQVHLCKE"
    return tabla[dni%23]
    ### END SOLUTION

```

```python
# Escribe aquí tus pruebas
letra(31415927)
```

Output:
```
'C'
```

```python
def test_letra():
    
    base = 131232*23
    t = 'TRWAGMYFPDXBNJZSQVHLCKE'
    for i in range(23):
        dni = base + i
        l = letra(dni)
        print(f'{dni}, {l}.....', end='')
        assert l == t[i], f'Error en dni {dni}, obtenido {l}, esperado {t[i]}'
        print('ok')
    print('ok')
    
test_letra()
```

Realiza ahora una función `checkdni_int` que compruebe el un DNI completo. La función recibe 2 parámetros: el dni y la letra. Debe comprobar que sea correcto

<!-- grade_id: cell-f2e0d406fd555b6b, points: 0, grade: false, solution: true -->
```python
def checkdni_int(dni:int, ltr: str) -> bool:
    ### BEGIN SOLUTION
    return letra(dni) == ltr
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-99785041350e6271, points: 1, grade: true, solution: false -->
```python
def test_checkdni_int():
    
    base = 131232*23
    t = 'TRWAGMYFPDXBNJZSQVHLCKE'
    for i in range(23):
        dni = base + i
        l = t[i]
        print(f'{dni}, {l}, True.....', end='')
        assert checkdni_int(dni, l), f'Error en dni {dni}, {l}'
        print('ok')
    print('ok')
    
    base = 131232*23 + 3 
    t = 'TRWAGMYFPDXBNJZSQVHLCKE'
    for i in range(23):
        dni = base + i
        l = t[i]
        print(f'{dni}, {l}, False.....', end='')
        assert not checkdni_int(dni, l), f'Error en dni {dni}, {l}'
        print('ok')
    print('ok')
    
test_checkdni_int()
```

Output:
```
3018336, T, True.....ok
3018337, R, True.....ok
3018338, W, True.....ok
3018339, A, True.....ok
3018340, G, True.....ok
3018341, M, True.....ok
3018342, Y, True.....ok
3018343, F, True.....ok
3018344, P, True.....ok
3018345, D, True.....ok
3018346, X, True.....ok
3018347, B, True.....ok
3018348, N, True.....ok
3018349, J, True.....ok
3018350, Z, True.....ok
3018351, S, True.....ok
3018352, Q, True.....ok
3018353, V, True.....ok
3018354, H, True.....ok
3018355, L, True.....ok
3018356, C, True.....ok
3018357, K, True.....ok
3018358, E, True.....ok
ok
3018339, T, False.....ok
3018340, R, False.....ok
3018341, W, False.....ok
3018342, A, False.....ok
3018343, G, False.....ok
3018344, M, False.....ok
3018345, Y, False.....ok
3018346, F, False.....ok
3018347, P, False.....ok
3018348, D, False.....ok
3018349, X, False.....ok
3018350, B, False.....ok
3018351, N, False.....ok
3018352, J, False.....ok
3018353, Z, False.....ok
3018354, S, False.....ok
3018355, Q, False.....ok
3018356, V, False...
```

Finalmente hay que comprobar si una cadena de caracteres contiene un DNI correcto. Supndremos que la parte inicial, desde la primera posición hasta la penúltima contiene dígitos decimales y la última posición contiene una letra en mayúsculas. Realiza una función que se llame `checkdni`

<!-- grade_id: cell-7425d7018d11fb32, points: 0, grade: false, solution: true -->
```python
def checkdni(dni: str) -> bool:
    ### BEGIN SOLUTION
    return checkdni_int(int(dni[:-1]), dni[-1]) 
    ### END SOLUTION
```

```python
# Escribe aquí tus pruebas
```

```python
def test_checkdni():
    
    base = 131232*23
    t = 'TRWAGMYFPDXBNJZSQVHLCKE'
    for i in range(23):
        dni = f'{base + i}{t[i]}'
        print(f'{dni}, True.....', end='')
        assert checkdni(dni), f'Error en dni {dni}, {l}'
        print('ok')
    print('ok')
    
    base = 131232*23 + 3 
    t = 'TRWAGMYFPDXBNJZSQVHLCKE'
    for i in range(23):
        dni = f'{base + i}{t[i]}'
        print(f'{dni}, False.....', end='')
        assert not checkdni(dni), f'Error en dni {dni}, {l}'
        print('ok')
    print('ok')
    
test_checkdni()    
```

Output:
```
3018336T, True.....ok
3018337R, True.....ok
3018338W, True.....ok
3018339A, True.....ok
3018340G, True.....ok
3018341M, True.....ok
3018342Y, True.....ok
3018343F, True.....ok
3018344P, True.....ok
3018345D, True.....ok
3018346X, True.....ok
3018347B, True.....ok
3018348N, True.....ok
3018349J, True.....ok
3018350Z, True.....ok
3018351S, True.....ok
3018352Q, True.....ok
3018353V, True.....ok
3018354H, True.....ok
3018355L, True.....ok
3018356C, True.....ok
3018357K, True.....ok
3018358E, True.....ok
ok
3018339T, False.....ok
3018340R, False.....ok
3018341W, False.....ok
3018342A, False.....ok
3018343G, False.....ok
3018344M, False.....ok
3018345Y, False.....ok
3018346F, False.....ok
3018347P, False.....ok
3018348D, False.....ok
3018349X, False.....ok
3018350B, False.....ok
3018351N, False.....ok
3018352J, False.....ok
3018353Z, False.....ok
3018354S, False.....ok
3018355Q, False.....ok
3018356V, False.....ok
3018357H, False.....ok
3018358L, False.....ok
3018359C, False.....ok
3018360K
```