# Cambio moneda

Supongamos que tenemos una máquina expendedora capaz de devolver cambio. Para ello dispone de monedas de 1, 5, 10 y 50 céntimos. 
Debemos de realizar una función que dada una cantidad de dinero, expresada en céntimos de euro, nos indique el número mínimo de monedas que debemos devolver de cada tipo. La función deberá devolver una terna de valores `(monedas de 50, monedas de 10, monedas de 1)`


<!-- grade_id: cell-811897d1db678576, points: 0, grade: false, solution: true -->
```python
def change(amount: int) -> (int, int, int):
    ### BEGIN SOLUTION
    fifties = amount // 50
    amount = amount % 50
    tens = amount // 10
    amount = amount % 10
    fives = amount // 5
    amount = amount % 5
    ### END SOLUTION
    return (fifties, tens, fives, amount)
```

```python
### Escribe aquí tus pruebas
change(78)
```

Output:
```
(1, 2, 1, 3)
```

<!-- grade_id: cell-759cefd3c3ff9d32, points: 0, grade: true, solution: false -->
```python
def test_change():
    for i in range(100):
        (fifties, tens, fives, ones) = change(i)
        assert i == fifties*50 + tens*10 + fives*5 + 1*ones, f'Error {i}, 50: {fifties}, 10: {tens}, 5: {fives}, 1:{ones}'
    print('OK')
test_change()
```

Output:
```
OK
```