# Años bisiestos

El antiguo calendario Juliano (introducido por Julio César) [ https://es.wikipedia.org/wiki/Calendario_juliano ] establecía un año bisiesto cada 4 años. Los cual asumía que un año tiene 365,25 días.
El error introducido provocó que el año 1582, el equinocio de primavera fuera el 11 de marzo en lugar de 21 de marzo. Puesto que esta fecha es clave en el calendario cristiando porque determina la fecha en la que se debe celebrar la Semana Santa, el papa Gregorio XIII creo la comsión del Calendario en la que destacó la figura del matemático y astónomo español Cristóbal Clavio. Esta comisión determinó la duración del año en 365,2422 días. La nueva norma de los años bisiestos se formuló del siguiente modo: la duración básica del año es de 365 días; pero serán bisiestos (es decir tendrán 366 días) aquellos años divisibles por 4, exceptuando los múltiplos de 100 (1700, 1800, 1900..., que no serán bisiestos), de los que se exceptúan a su vez aquellos que también sean divisibles por 400 (1600, 2000, 2400..., que serán bisiestos). El calendario gregoriano ajusta a 365,2425 días la duración del año, lo que deja una diferencia de 0,0003 días al año de error, es decir, adelanta cerca de 1/2 minuto cada año (aprox. 26 s/año), lo que significa que se requiere el ajuste de un día cada 3323 años.

Para ajustar el calendario, al jueves (juliano) 4 de octubre de 1582 le sucedió el viernes (gregoriano) 15 de octubre de 1582. Así, diez días desaparecieron debido a que ya se habían contado de más en el calendario juliano. 

Extracto de la wikipedia [ https://es.wikipedia.org/wiki/Calendario_gregoriano ]

Haz una función `is_leap_year` que indique si un año mayor que 1582 es bisiesto o no 

<!-- grade_id: cell-53807fab75031d29, points: 0, grade: false, solution: true -->
```python
def is_leap_year(year: int) -> bool:
    ### BEGIN SOLUTION
    return year % 4 == 0 and (year % 100 !=0 or year % 400 == 0)
    ### END SOLUTION
```

```python
### Escribe tus pruebas
```

<!-- grade_id: cell-9d5d5d1535db056e, points: 1, grade: true, solution: false -->
```python
def test_is_leap_year():
    tests = [ (1583, False), (1600, True), (1601, False), (1700, False), (2000, True), (2001, False), (2004, True), (2100, False) ]
    for year, is_leap in tests:
        res = is_leap_year(year)
        assert res == is_leap, f'Error {year}, {is_leap}, result: {res}'
    print('OK')
    
test_is_leap_year()
```

Output:
```
OK
```