# Recorridos y búsquedas en cadenas de caracteres

En estas funciones sólo se pueden realizar accesos a las componentes de la cadena de caracteres `s[pos]`, función `len`, y concatenación de cadenas de caracteres (operador `+`).

Realiza una función que nos indique cuántas apariciones de una letra hay en una cadena de caracteres.

<!-- grade_id: cell-cf5326ab86db4507, points: 0, grade: false, solution: true -->
```python
def count_occurrences(s: str, ltr:str) -> int:
    ### BEGIN SOLUTION
    assert len(ltr) == 1
    i = 0
    count = 0
    while i<len(s):
        if s[i] == ltr:
            count = count + 1
        i = i + 1
    return count
    ### END SOLUTION
```

```python

```

<!-- grade_id: cell-5097dda96cae9762, points: 1, grade: true, solution: false -->
```python
from functools import total_ordering

@total_ordering

class WrapperStr(str):
    def __init__(self, s: str):
        self.str = s
    def __eq__(self, other):
            if len(self.str)!=1 or len(other)!=1:
                raise Exception(f'It is illegal to compare strings of lengh greater than 1: "{self.str}", "{other.str}"')
            return self.str==other
    def __le__(self, other):
            if len(self.str)!=1 or len(other)!=1:
                raise Exception(f'It is illegal to compare strings of lengh greater than 1: "{self.str}", "{other.str}"')
            return self.str<other
    
    def __contains__(self, item):
            raise Exception(f'Illegal call to "in"')
    def __getitem__(self, key):
        if not isinstance(key, int):
            raise Exception(f'Illegal call [{key}]')
        return self.str[key]
    
def runtests(tests, fun):
    for test, ex_res in tests:
        print(f'"{test}", "{ex_res}"....', end='')
        res = fun(*test)
        assert res == ex_res, f'Error, the result is "{res}"'
        print('OK.')

        
def test_count_occurrences():
    def wrapper(s: str, ltr: str):
        return count_occurrences(WrapperStr(s), WrapperStr(ltr))
    
    tests = [
        (('', 'g'),0),        
        (('abc', 'g'),0),
        (('abaca', 'a'),3),
        (('abaca', 'b'),1),
    ]
    runtests(tests, wrapper)
    
    try:
        s = count_occurrences('a', 'bc')
    except AssertionError:
        pass
    else:
        assert False, f'The function must check that lentgh of the second argument is 1'
    
test_count_occurrences()
```

Output:
```
"('', 'g')", "0"....OK.
"('abc', 'g')", "0"....OK.
"('abaca', 'a')", "3"....OK.
"('abaca', 'b')", "1"....OK.
```

Realiza una función que indique si una letra aparece en una secuencia de caracteres

```python
### Función no válida
def is_in(s: str, ltr: str) -> bool:
    return ltr in s
```

<!-- grade_id: cell-f5592fa560c7e7c5, points: 0, grade: false, solution: true -->
```python
def is_in(s: str, ltr:str) -> bool:
    ### BEGIN SOLUTION
    i = 0
    while i<len(s) and s[i]!=ltr:
        i = i + 1
    return i<len(s)
    ### END SOLUTION
```

```python

```

<!-- grade_id: cell-fcaadffb75eea7d2, points: 1, grade: true, solution: false -->
```python
def test_is_in():
    def wrapper(s1: str, s2:str):
        return is_in(WrapperStr(s1), WrapperStr(s2))

    tests = [
        (('', 'g'),False),        
        (('abc', 'g'),False),
        (('abaca', 'a'),True),
        (('abaca', 'b'),True),
        (('abac', 'c'),True),
    ]
    runtests(tests, wrapper)
    
test_is_in()
```

Output:
```
"('', 'g')", "False"....OK.
"('abc', 'g')", "False"....OK.
"('abaca', 'a')", "True"....OK.
"('abaca', 'b')", "True"....OK.
"('abac', 'c')", "True"....OK.
```

Realiza una función que compruebe si dos cadenas de caracteres son iguales. Obviamente no vale hacer `return s1==s2`. Solo se debe comprobar la igualdad entre posiciones de las cadenas de caracteres del tipo `s2[1]==s1[1]`. La función se debe llamar `are_equal`

```python
### Esta función no vale
def are_equal(s1: str, s2:str) -> bool:
    return s1 == s2
```

<!-- grade_id: cell-b3341994899df8fc, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def are_equal(s1: str, s2:str) -> bool:
    res = False
    if len(s1)==len(s2):
        i = 0
        while i<len(s1) and s1[i]==s2[i]:
            i = i + 1
        res = i==len(s1)
    return res
### END SOLUTION
```

```python

```

<!-- grade_id: cell-6a5207d422a667c4, points: 1, grade: true, solution: false -->
```python

def test_are_equal():
    def wrapper(s1: str, s2:str):
        return are_equal(WrapperStr(s1), WrapperStr(s2))
                                
    tests = [
        (('', 'g'),False),        
        (('abc', 'g'),False),
        (('g', ''),False),
        (('aba', 'ab'),False),
        (('ab', 'aba'),False),
        (('', ''),True),        
        (('abc', 'abc'),True),        
    ]
    runtests(tests, wrapper)

test_are_equal()
```

Output:
```
"('', 'g')", "False"....OK.
"('abc', 'g')", "False"....OK.
"('g', '')", "False"....OK.
"('aba', 'ab')", "False"....OK.
"('ab', 'aba')", "False"....OK.
"('', '')", "True"....OK.
"('abc', 'abc')", "True"....OK.
```

Realiza una función que admita 2 cadenas de caractes e indique si la primera es anterior a la segunda según el orden lexicográfico. Al igual que en el ejercicio anterior solo se pueden hacer comparaciones de una sola letra. La función se debe llamar `less_than`

```python
### Esta función no vale
def less_than(s1: str, s2:str) -> bool:
    return s1 == s2
```

<!-- grade_id: cell-f9239a91e0565956, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def less_than(s1: str, s2:str) -> bool:
    i = 0
    while i<len(s1) and i<len(s2) and s1[i]==s2[i]:
        i = i + 1
        
    if i < len(s2) and i==len(s1):
        res = True
    elif i<=len(s1) and i==len(s2):
        res = False
    else:
        res = s1[i]<s2[i]
    return res
### END SOLUTION
```

<!-- grade_id: cell-caa0ae6167cd54af, points: 1, grade: true, solution: false -->
```python
                                
def test_less_than():
    def wrapper(s1: str, s2:str):
        return less_than(WrapperStr(s1), WrapperStr(s2))

    tests = [
        (('', 'g'),True),        
        (('abc', 'g'),True),
        (('g', ''),False),
        (('aba', 'ab'),False),
        (('ab', 'aba'),True),
        (('', ''),False),        
        (('abc', 'abc'),False),        
        (('abc', 'abd'),True),        
    ]
    runtests(tests, wrapper)


test_less_than()
```

Output:
```
"('', 'g')", "True"....OK.
"('abc', 'g')", "True"....OK.
"('g', '')", "False"....OK.
"('aba', 'ab')", "False"....OK.
"('ab', 'aba')", "True"....OK.
"('', '')", "False"....OK.
"('abc', 'abc')", "False"....OK.
"('abc', 'abd')", "True"....OK.
```

Implementar una función que se llame `slice` que admita como parámteros una cadena de caracteres, una posición inicial y una posición final. La función debe devolver los caracteres entre la posición inicial (incluida) y la posición final (sin incluir). Se debe comprobar que ambos índices toman valores no negativos y que estén en las posiciones adecuadas.

```python
### Función no válida
def slice(s: str, ini:int, end: int) -> str:
    return s[ini:end]
```

<!-- grade_id: cell-ca0a43fa04621ec7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def slice(s: str, ini:int, end: int) -> str:
    assert 0<=ini<len(s), f'The initial position must be a legal position'
    assert 0<=end<=len(s), f'The final position must be a legal position'
    i = ini
    cad = ''
    while i<end:
        cad = cad + s[i]
        i = i+1
    return cad
### END SOLUTION
    
```

```python

```

<!-- grade_id: cell-ab1fd055e866469d, points: 0, grade: true, solution: false -->
```python
                                
def test_less_than():
    def wrapper(s1: str, ini:int, end: int) -> str:
        return slice(WrapperStr(s1), ini, end)

    tests = [
        (('abc', 0, 0), ''),
        (('abc', 0, 1), 'a'),
        (('abc', 0, 3), 'abc'),
        (('abc', 1, 3), 'bc'),
        (('abc', 2, 3), 'c'),
        (('abc', 2, 1), ''),
    ]
    runtests(tests, wrapper)
    
    try:
        s = slice('a', -1, 1)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the initial position must be non-negative'
        
    try:
        s = slice('a', 1, -1)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the final position must be non-negative'

    try:    
        s = slice('a', 2, 2)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the inital position is a legal position'

    try:    
        s = slice('a', 0, 2)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the final position is a legal position'



test_less_than()
```

Ahora vamos a hacer una función que nos devuelva la siguiente aparición de una letra en la cadena de caracteres. La fución se deberá llamar `next_occurrence`, y debe admitir como parámetros (en este orden) una cadena de caractares donde vamos a buscar, un carácter y la posición inicial donde queremos buscar. La función debe
* Devolver la siguente posición en la que aparece la letra que buscamos, que en particular puede ser la posición que en la que nos encontramos. 
* Si la letra no aparece se devuelve la longitud de la cadena de caracteres.
* La función debe comprobar que la posición inicial es una posición 'legal' de la cadena de caracteres. 

<!-- grade_id: cell-a14fa2f334a84703, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def next_occurrence(s: str, ltr:str, pos:int) -> int:
    assert len(ltr) == 1
    assert pos>=0, 'The initial position must be non negative'
    assert pos<len(s), f'The initial position must be less than len({s}): {pos}'
    i = pos
    while i<len(s) and s[i]!=ltr:
        i = i + 1
    return i
### END SOLUTION
                    
```

```python

```

<!-- grade_id: cell-cefb77318c631a38, points: 1, grade: true, solution: false -->
```python
                                
def test_next_ocurrence():
    def wrapper(s: str, ltr:str, pos: int) -> int:
        return next_occurrence(WrapperStr(s), ltr, pos)
                                
    tests = [
        (('abc', 'a', 0), 0),
        (('abc', 'c', 1), 2),
        (('abc', 'b', 2), 3),
    ]
    runtests(tests, wrapper)
    
    try:
        s = next_occurrence('a', 'bc', 0)
    except AssertionError:
        pass
    else:
        assert False, f'The function must check that lentgh of the second argument is 1'
    
    try:
        s = next_occurrence('a', 'a', -1)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the inicial position must be non-negative'

    try:    
        s = next_occurrence('a', 'a', 2)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the inital position is a legal position'    
    
test_next_ocurrence()
```

Realiza una función que indique si una subcadena está contendida en otra. La función se debe llamar `contains`. El primer argumento es la cadena donde buscamos y el segundo argumento es la cadena buscada. Ejemplos:

        contains('abc', 'a') ---> True
        contains('abc', 'c') --->  True
        contains('abc', 'bc') ---> True
        contains('abc', '') ---> True
        contains('abc', 'd') ---> False
        
**PISTA:** Usa las funciones `slice` y `are_equal` definidas anteriormente.

<!-- grade_id: cell-4348be59b2d313a9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def contains(big: str, small: str) -> bool:
    i = 0
    while i < len(big)-len(small)+1 and not are_equal(slice(big, i, i+len(small)), small):
        i = i +1
    return i < len(big)-len(small)+1

### END SOLUTION
```

```python

```

<!-- grade_id: cell-87a4f685e66db221, points: 1, grade: true, solution: false -->
```python
def test_contains():
    def wrapper(s1: str, s2:str) -> bool:
        return contains(WrapperStr(s1), WrapperStr(s2))

    tests = [
        (('abc', 'a'), True),
        (('abc', 'c'), True),
        (('abc', 'bc'), True),
        (('abc', ''), True),
        (('abc', 'd'), False),
        (('abc', 'dada'), False),
    ]
    runtests(tests, wrapper)

test_contains()
```

Ahora vamos a hacer una función parecida a la que teníamos de next_ocurrence, pero ahora vamos quitar la restricción que que el segundo sea un parámetro y sea una cadena de caracteres arbitraria. 
Si la cadena no está devolvemos la longitud de la primera cadena. La función se debe llamar `next_occurrence`

Ejemplos

        next_occurence('ababc', 'ab', 1) ---> 2
        next_occurence('ababc', 'cd', 1) ---> 5
        

<!-- grade_id: cell-7524f77203423f14, points: 0, grade: false, solution: true -->
```python
### Este código es para evitar errores por la función anterior
try:
    del next_occurrence
except:
    pass
##############################################################    
    
### BEGIN SOLUTION
def next_occurrence(big: str, small: str, pos: int) -> int:
    assert pos>=0
    assert pos<len(big)
    i = pos
    while i < len(big)-len(small)+1 and not are_equal(slice(big, i, i+len(small)), small):
        i = i +1
    if i < len(big)-len(small)+1:
        res = i
    else:
        res = len(big)
    return res
        
### END SOLUTION
```

```python

```

<!-- grade_id: cell-9ec39779a2acbd54, points: 1, grade: true, solution: false -->
```python
def test_next_ocurrence():
    def wrapper(s1: str, s2:str, pos: int) -> int:
        return next_occurrence(WrapperStr(s1), WrapperStr(s2), pos)
                                
    tests = [
        (('ababcaba', 'ab', 0), 0),
        (('ababcaba', 'ab', 1), 2),
        (('ababcaba', 'ab', 2), 2),
        (('ababcaba', 'ab', 3), 5),
        (('ababcaba', 'ab', 6), 8),
    ]
    runtests(tests, wrapper)
        
    try:
        s = next_occurrence('a', 'a', -1)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the inicial position must be non-negative'

    try:    
        s = next_occurrence('a', 'a', 2)
    except AssertionError:
        pass
    else:
        assert False, 'The function does not check that the inital position is a legal position'    
    
test_next_ocurrence()
```