# Fecha absoluta

Con todo esto, representar una fecha como se hace habitualmente
`(dia,mes,anyo)` no parece muy conveniente. La forma habitual 
de hacerlo en una computadora es contar los milisegundos desde
una fecha previamente establecida. Nosotros lo vamos a
simplificar un poco, en primer lugar como sólo queremos
establecer la fecha (no la hora), vamos a contar los días
transcurridos desde una fecha prefijada. 

En segundo lugar, 
para facilitar los
cómputos estableceremos la fecha inicial el 1 de enero de
1601, este será nuestro día 0. Este día marca el comienzo de un bloque de 400 años que
culmina en el año bisiesto 2000. De esta forma, 

* El 23 de abril de 1616 es el día 5591.
* El 10 de octubre de 1968 es el día  134326.
* El 1 de enero de 2001 (el primer día de este siglo) es el
día 146097.
* El 27 de noviembre de 2006 es el día 148253.

Por último como sabemos que el 1 de enero de 1601 
fue lunes,
podemos saber qué día de la semana es cualquier fecha en el
calendario gregoriano. Así sabemos que a Cervantes lo enterraron un 
sábado [ https://es.wikipedia.org/wiki/Miguel_de_Cervantes#Muerte_y_tumba_de_Cervantes ]. 

Recordemos que un año es bisiesto si es múltiplo de 4,
salvo si es múltiplo de 100 en cuyo caso debe ser también
múltiplo de 400.

**RECUERDA**: Las funciones deben tener la documentación de forma correcta y deben capturar los errores previsibles.


## Anyo bisiesto
Haz una función `is_leap_year` que indique si un año mayor que 1601 es bisiesto o no 

<!-- grade_id: cell-dbabb54c91356c64, points: 0, grade: false, solution: true -->
```python
### 1 PUNTO

def is_leap_year(year: int) -> bool:
    ### BEGIN SOLUTION
    assert year>1600, f'The year must be greater than 1600'
    return year % 4 == 0  and ((year % 100 != 0) or (year % 400 == 0))
    ### END SOLUTION
```

<!-- grade_id: cell-0f5ac651718bd89b, points: 1, grade: true, solution: false -->
```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-fef1f63091bfef57, points: 1, grade: true, solution: false -->
```python
def test_is_leap_year():
    tests = [ (1601, False),
              (1700, False),
              (2000, True),
              (2001, False),
              (2004, True),
              (2100, False)
    ]

    for year, is_leap in tests:
        res = is_leap_year(year)
        assert res == is_leap, f'Error {year}, {is_leap}, result: {res}'

    for i in range(1582, 1000, -1):
        try:
            res = is_leap_year(1500)
        except AssertionError:
            pass
        else:
            assert False, f'No se ha lanzado el error para un año {i} menor que 1582'
    print('OK')

test_is_leap_year()

```

Output:
```
OK
```

## Días mes

Haz una función `days_month` que devuelve el número de días que tiene un mes. Como el número de días de febrero depende de si el año es bisiesto o no, la función debe tener como parámetro un `bool` que indique si el anyo es bisiesto o no. El mes es un entero entre 1 y 12, el año es un entero mayor que 1582. 

<!-- grade_id: cell-90eeee25629cce42, points: 0, grade: false, solution: true -->
```python
### 2 PUNTOS

def days_month(month:int, year:int) -> bool:
    ### BEGIN SOLUTION
    assert year> 1600, f'The year must be greater than 1600'
    assert 1<=month and month<=12, f'The month must be an integer bwtween 1 and 12'
    if month == 2:
        days = 28
        if is_leap_year(year):
            days = days + 1
    elif (month % 2 == 1 and month <= 7) or \
       (month % 2 == 0 and month >= 8):
       days = 31
    else:
        days = 30
    return days
    ### END SOLUTION

```

```python
# Escribe aquí tus pruebas
```

<!-- grade_id: cell-4101f05f67189b24, points: 2, grade: true, solution: false -->
```python
def test_days_month():
    tests = [
        ((1,2001),31),
        ((2,2001),28),
        ((3,2001),31),
        ((4,2001),30),
        ((5,2001),31),
        ((6,2001),30),
        ((7,2001),31),
        ((8,2001),31),
        ((9,2001),30),
        ((10,2001),31),
        ((11,2001),30),
        ((12,2001),31),
        ((2,2000),29),
    ]
    for val, ex_res in tests:
        res = days_month(*val)
        print(f'{val} {res}...', end='')
        assert ex_res == res, f'the expected value is {ex_res}'
        print('OK')

    try:
        res = days_month(0,2000)
    except AssertionError:
        pass
    else:
        assert False, f'The function should detect that the month is incorrect'

    try:
        res = days_month(13,2000)
    except AssertionError:
        pass
    else:
        assert False, f'The function should detect that the month is incorrect'

    try:
        res = days_month(1,1600)
    except AssertionError:
        pass
    else:
        assert False, f'The function should detect that the year is incorrect'


    try:
        res = days_month(1,1500)
    except AssertionError:
        pass
    else:
        assert False, f'The function should detect that the year is incorrect'
    print('OK')
test_days_month()

```

Output:
```
(1, 2001) 31...OK
(2, 2001) 28...OK
(3, 2001) 31...OK
(4, 2001) 30...OK
(5, 2001) 31...OK
(6, 2001) 30...OK
(7, 2001) 31...OK
(8, 2001) 31...OK
(9, 2001) 30...OK
(10, 2001) 31...OK
(11, 2001) 30...OK
(12, 2001) 31...OK
(2, 2000) 29...OK
OK
```

## Fecha correcta

Realiza una función `is_correct_date` que indique si una terna `(dia, mes, anyo)` representa una fecha correcta. Es decir, que compruebe que los elementos son correctos:
* El año debe ser mayor estricto que 1600.
* El mes debe ser un entero entre 1 y 12.
* El día debe ser un entero entre 1 y el número de días que tenga el mes correspondiente. Ten en cuenta que, en los años bisistos, febrero tiene 29 días.

<!-- grade_id: cell-6b31ab529c67804f, points: 0, grade: false, solution: true -->
```python
### 1 PUNTO

def is_correct(day, month, year):
    ### BEGIN SOLUTION
    """
    This function returns True if day/month/year respresents
    a valid date after 1/1/1601
    """
    return year >= 1601 and \
        (1 <= month <= 12) and \
        (1 <= day <= days_month(month, year))
    ### END SOLUTION

```

```python

```

<!-- grade_id: cell-391e5be53d2c6331, points: 1, grade: true, solution: false -->
```python
def test_is_correct():
    tests = [
        ((1,1,1601), True),
        ((0,1,1601), False),
        ((32,1,1601), False),
        ((1,0,1601), False),
        ((1,13,1601), False),
        ((31,1,1601), True),
        ((28,2,1601), True),
        ((31,3,1601), True),
        ((30,4,1601), True),
        ((31,5,1601), True),
        ((30,6,1601), True),
        ((31,7,1601), True),
        ((31,8,1601), True),
        ((30,9,1601), True),
        ((31,10,1601), True),
        ((30,11,1601), True),
        ((31,12,1601), True),
        ((29,2,2000), True),
        ((29,2,1900), False),
    ]
    for val, ex_res in tests:
        res = is_correct(*val)
        print(f'{val} {res}...', end='')
        assert ex_res == res, f'the expected value is {ex_res}'
        print('OK')
test_is_correct()

```

Output:
```
(1, 1, 1601) True...OK
(0, 1, 1601) False...OK
(32, 1, 1601) False...OK
(1, 0, 1601) False...OK
(1, 13, 1601) False...OK
(31, 1, 1601) True...OK
(28, 2, 1601) True...OK
(31, 3, 1601) True...OK
(30, 4, 1601) True...OK
(31, 5, 1601) True...OK
(30, 6, 1601) True...OK
(31, 7, 1601) True...OK
(31, 8, 1601) True...OK
(30, 9, 1601) True...OK
(31, 10, 1601) True...OK
(30, 11, 1601) True...OK
(31, 12, 1601) True...OK
(29, 2, 2000) True...OK
(29, 2, 1900) False...OK
```

## Días desde el comienzo

Vamos a hacer una función que calcule el número de días que han transcurrido desde el 1/1/1601 hasta una fecha dada. La función debe comprobar la corrección de la fecha introducida como parámetro.

<!-- grade_id: cell-d3b3afb3be808bf1, points: 0, grade: false, solution: true -->
```python
### 3 PUNTOS


INIT_YEAR = 1601

def count_days_from_init(day, month, year):
    ### BEGIN SOLUTION
    """
    This function computes the number of days that has passed
    since 1/1/1601
    @rtype int
    @return: >=0
    @type day: int
    @type month: int
    @type year: int
    @param day,month,year: they represent a valid date after 1/1/1601
    """

    assert is_correct(day, month, year), f'Error fecha incorrecta: {day}/{month}/{year}'
    num_years = year - INIT_YEAR
    num_years4 = num_years // 4
    num_years100 = num_years // 100
    num_years400 = num_years // 400

    # days until 1/1/year (excluding it)
    num_days = num_years * 365 + num_years4 - num_years100 + num_years400


    # num days until 1/month/year asuming all months
    # have 30 days
    num_days += (month-1)*30

    # we adjust for the months that do not have 30 years
    if month > 2:
        if is_leap_year(year):
            num_days -= 1
        else:
            num_days -= 2

    if month <= 8:
        num_days += month//2
    else:
        num_days += 4 + (month-8+1)//2

    num_days += day - 1

    return num_days
    ### END SOLUTION

```

```python

```

<!-- grade_id: cell-cfaf5948a466b090, points: 3, grade: true, solution: false -->
```python
def test_count_days_from_init():
    tests = [
        (0,1,1601),
        (32,1,1601),
        (1,0,1601),
        (1,13,1601),
        (29,2,1900),
    ]
    for date in tests:
        try:
            n = count_days_from_init(*date)
        except AssertionError:
            pass
        else:
            assert False, f'The function does not check the validity of the dates'
    print('OK')


    tests = [
        ((1,1,1601), 0),
        ((31,1,1601), 30),
        ((1,1,1602), 365),
        ((1,2,1602), 365+31),
        ((29,2,1604), 365*3+31+28),
        ((3,3,1604), 365*3+31+29+2),
        ((23,4,1616), 5591),
        ((10,10,1968), 134326),
        ((1,1,2001), 146097),
        ((27,11,2006), 148253)
    ]
    for val, ex_res in tests:
        res = count_days_from_init(*val)
        print(f'{val} {res}...', end='')
        assert ex_res == res, f'the expected value is {ex_res}'
        print('OK')
        
test_count_days_from_init()

```

Output:
```
OK
(1, 1, 1601) 0...OK
(31, 1, 1601) 30...OK
(1, 1, 1602) 365...OK
(1, 2, 1602) 396...OK
(29, 2, 1604) 1154...OK
(3, 3, 1604) 1157...OK
(23, 4, 1616) 5591...OK
(10, 10, 1968) 134326...OK
(1, 1, 2001) 146097...OK
(27, 11, 2006) 148253...OK
```

Ahora hay que hacer la función inversa. Dado un entero $n\geq 0$ hay que averguar a qué fecha se corresponde. Recuerda que empezamos a contar en el día 0, que se corresponde con el 1/1/1601, y por tanto el 364 será el 31/12/1601. Puedes usar la siguiente función, que no está demasiado bien diseñada.

```python
def get_day_month(ndays: int, year:int)-> (int, int):
    '''
    This function returns the day and month that corresponds to an amount of
    ndays in the year

    Precondition
    ------------
    if year is a leap year, 0<=ndays<366
    otherwise,  0<=ndays<366
    

    Examples:
    --------
In [232]: get_day_month(0, 2000)
Out[240]: (1, 1)

In [241]: get_day_month(31, 2000)
Out[241]: (1, 2)
    
In [242]: get_day_month(365, 2000)
Out[242]: (31, 12)

In [243]: get_day_month(364, 2001)
Out[243]: (31, 12)

    '''
    assert ndays>=0 and (ndays<366 and is_leap_year(year) or \
        ndays<365, f'The thear {year} has less that {ndays} days')
    
    day = ndays + 1 # we start counting on day 0,
    #but months start countin on 1

    """Esto que sigue es el típico código spaghetti,
    ¡¡¡una verdadera chapuza!!!.
    Hay que hacerlo con bucles."""

    month = 1
    if day > days_month(month, year): #No es enero
        day -= days_month(month, year)
        month = month + 1
        if day > days_month(month, year): # no es febrero
            day -= days_month(month, year)
            month = month + 1
            if day > days_month(month, year): # no es marzo
                day -= days_month(month, year)
                month = month + 1
                if day > days_month(month, year): #no es abril
                    day -= days_month(month, year)
                    month = month + 1
                    if day > days_month(month, year): #no es mayo
                        day -= days_month(month, year)
                        month = month + 1
                        if day > days_month(month, year): #no es junio
                            day -= days_month(month, year)
                            month = month + 1
                            if day > days_month(month, year): #no es julio
                                day -= days_month(month, year)
                                month = month + 1
                                if day > days_month(month, year): #no es agosto
                                    day -= days_month(month, year)
                                    month = month + 1
                                    if day > days_month(month, year): #no es septiembre
                                        day -= days_month(month, year)
                                        month = month + 1
                                        if day > days_month(month, year): #no es octubre
                                            day -= days_month(month, year)
                                            month = month + 1
                                            if day > days_month(month, year): #no es noviembre
                                                day -= days_month(month, year)
                                                month = month + 1
                                                if day > days_month(month, year): #no es diciembre
                                                    raise Exception(f'Error internor, datos imposibles: {day} {month}')
    return day, month

```

```python
def test_get_day_month():
    # for i in range(366):
    #     (d,m) = get_day_month(i, 2000)
    #     if d == 1:
    #         print('')
    #     if d in [1,2,28,29,30,31]:
    #         print(f'(({i}, 2000), ({d}, {m})), ', end='')

    # for i in range(365):
    #     (d,m) = get_day_month(i, 2001)
    #     if d == 1:
    #         print('')
    #     if d in [1,2,28,29,30,31]:
    #         print(f'(({i}, 2001), ({d}, {m})), ', end='')

    tests = [
        ((0, 2000), (1, 1)), ((1, 2000), (2, 1)), ((27, 2000), (28, 1)), ((28, 2000), (29, 1)), ((29, 2000), (30, 1)), ((30, 2000), (31, 1)),
        ((31, 2000), (1, 2)), ((32, 2000), (2, 2)), ((58, 2000), (28, 2)), ((59, 2000), (29, 2)),
        ((60, 2000), (1, 3)), ((61, 2000), (2, 3)), ((87, 2000), (28, 3)), ((88, 2000), (29, 3)), ((89, 2000), (30, 3)), ((90, 2000), (31, 3)),
        ((91, 2000), (1, 4)), ((92, 2000), (2, 4)), ((118, 2000), (28, 4)), ((119, 2000), (29, 4)), ((120, 2000), (30, 4)),
        ((121, 2000), (1, 5)), ((122, 2000), (2, 5)), ((148, 2000), (28, 5)), ((149, 2000), (29, 5)), ((150, 2000), (30, 5)), ((151, 2000), (31, 5)),
        ((152, 2000), (1, 6)), ((153, 2000), (2, 6)), ((179, 2000), (28, 6)), ((180, 2000), (29, 6)), ((181, 2000), (30, 6)),
        ((182, 2000), (1, 7)), ((183, 2000), (2, 7)), ((209, 2000), (28, 7)), ((210, 2000), (29, 7)), ((211, 2000), (30, 7)), ((212, 2000), (31, 7)),
        ((213, 2000), (1, 8)), ((214, 2000), (2, 8)), ((240, 2000), (28, 8)), ((241, 2000), (29, 8)), ((242, 2000), (30, 8)), ((243, 2000), (31, 8)),
        ((244, 2000), (1, 9)), ((245, 2000), (2, 9)), ((271, 2000), (28, 9)), ((272, 2000), (29, 9)), ((273, 2000), (30, 9)),
        ((274, 2000), (1, 10)), ((275, 2000), (2, 10)), ((301, 2000), (28, 10)), ((302, 2000), (29, 10)), ((303, 2000), (30, 10)), ((304, 2000), (31, 10)),
        ((305, 2000), (1, 11)), ((306, 2000), (2, 11)), ((332, 2000), (28, 11)), ((333, 2000), (29, 11)), ((334, 2000), (30, 11)),
        ((335, 2000), (1, 12)), ((336, 2000), (2, 12)), ((362, 2000), (28, 12)), ((363, 2000), (29, 12)), ((364, 2000), (30, 12)), ((365, 2000), (31, 12)),
        ((0, 2001), (1, 1)), ((1, 2001), (2, 1)), ((27, 2001), (28, 1)), ((28, 2001), (29, 1)), ((29, 2001), (30, 1)), ((30, 2001), (31, 1)),
        ((31, 2001), (1, 2)), ((32, 2001), (2, 2)), ((58, 2001), (28, 2)),
        ((59, 2001), (1, 3)), ((60, 2001), (2, 3)), ((86, 2001), (28, 3)), ((87, 2001), (29, 3)), ((88, 2001), (30, 3)), ((89, 2001), (31, 3)),
        ((90, 2001), (1, 4)), ((91, 2001), (2, 4)), ((117, 2001), (28, 4)), ((118, 2001), (29, 4)), ((119, 2001), (30, 4)),
        ((120, 2001), (1, 5)), ((121, 2001), (2, 5)), ((147, 2001), (28, 5)), ((148, 2001), (29, 5)), ((149, 2001), (30, 5)), ((150, 2001), (31, 5)),
        ((151, 2001), (1, 6)), ((152, 2001), (2, 6)), ((178, 2001), (28, 6)), ((179, 2001), (29, 6)), ((180, 2001), (30, 6)),
        ((181, 2001), (1, 7)), ((182, 2001), (2, 7)), ((208, 2001), (28, 7)), ((209, 2001), (29, 7)), ((210, 2001), (30, 7)), ((211, 2001), (31, 7)),
        ((212, 2001), (1, 8)), ((213, 2001), (2, 8)), ((239, 2001), (28, 8)), ((240, 2001), (29, 8)), ((241, 2001), (30, 8)), ((242, 2001), (31, 8)),
        ((243, 2001), (1, 9)), ((244, 2001), (2, 9)), ((270, 2001), (28, 9)), ((271, 2001), (29, 9)), ((272, 2001), (30, 9)),
        ((273, 2001), (1, 10)), ((274, 2001), (2, 10)), ((300, 2001), (28, 10)), ((301, 2001), (29, 10)), ((302, 2001), (30, 10)), ((303, 2001), (31, 10)),
        ((304, 2001), (1, 11)), ((305, 2001), (2, 11)), ((331, 2001), (28, 11)), ((332, 2001), (29, 11)), ((333, 2001), (30, 11)),
        ((334, 2001), (1, 12)), ((335, 2001), (2, 12)), ((361, 2001), (28, 12)), ((362, 2001), (29, 12)), ((363, 2001), (30, 12)), ((364, 2001), (31, 12)),
    ]
    for val, ex_res in tests:
        res = get_day_month(*val)
        print(f'{val} {res}...', end='')
        assert ex_res == res, f'the expected value is {ex_res}'
        print('OK')

test_get_day_month()
```

Output:
```
(0, 2000) (1, 1)...OK
(1, 2000) (2, 1)...OK
(27, 2000) (28, 1)...OK
(28, 2000) (29, 1)...OK
(29, 2000) (30, 1)...OK
(30, 2000) (31, 1)...OK
(31, 2000) (1, 2)...OK
(32, 2000) (2, 2)...OK
(58, 2000) (28, 2)...OK
(59, 2000) (29, 2)...OK
(60, 2000) (1, 3)...OK
(61, 2000) (2, 3)...OK
(87, 2000) (28, 3)...OK
(88, 2000) (29, 3)...OK
(89, 2000) (30, 3)...OK
(90, 2000) (31, 3)...OK
(91, 2000) (1, 4)...OK
(92, 2000) (2, 4)...OK
(118, 2000) (28, 4)...OK
(119, 2000) (29, 4)...OK
(120, 2000) (30, 4)...OK
(121, 2000) (1, 5)...OK
(122, 2000) (2, 5)...OK
(148, 2000) (28, 5)...OK
(149, 2000) (29, 5)...OK
(150, 2000) (30, 5)...OK
(151, 2000) (31, 5)...OK
(152, 2000) (1, 6)...OK
(153, 2000) (2, 6)...OK
(179, 2000) (28, 6)...OK
(180, 2000) (29, 6)...OK
(181, 2000) (30, 6)...OK
(182, 2000) (1, 7)...OK
(183, 2000) (2, 7)...OK
(209, 2000) (28, 7)...OK
(210, 2000) (29, 7)...OK
(211, 2000) (30, 7)...OK
(212, 2000) (31, 7)...OK
(213, 2000) (1, 8)...OK
(214, 2000) (2, 8)...OK
(240, 2000) (28, 8)...OK
(241, 2000)
```

<!-- grade_id: cell-a07fe298664a1b96, points: 0, grade: false, solution: true -->
```python
# 4 PUNTOS 


BLOCK_400 = 365*400 + 100 - 4 + 1
BLOCK_100 = 365*100 + 25 - 1
BLOCK_4 = 365*4 +  1

def gregorian_date(ndays):
    ### BEGIN SOLUTION
    """
    Function that computes the day/month/year with respect
    the gregorian calendar
    @rtype: 3-integer tuple (day,month/myear)
    """
    assert ndays>=0, f'The ndays cannot be negative: {ndays}'

    year = INIT_YEAR
    year += 400 * (ndays//BLOCK_400)
    ndays = ndays % BLOCK_400

    if ndays == BLOCK_400 - 1:
        year += 300
        ndays = BLOCK_100
    else:
        year += 100 * (ndays//BLOCK_100)
        ndays = ndays % BLOCK_100

    year += 4 * (ndays//BLOCK_4)
    ndays = ndays % BLOCK_4

    if ndays == BLOCK_4 - 1:
        year += 3
        ndays = 365
    else:
        year += ndays // 365
        ndays = ndays % 365

    day, month = get_day_month(ndays, year)
    return day, month, year
    ### END SOLUTION
    
    
```

```python
### Escribe tus pruebas
```

<!-- grade_id: cell-9a4435d36572f6a9, points: 4, grade: true, solution: false -->
```python
def test_gregorian_date():
    import random
    tests = [
        (0, (1, 1, 1601)),
        (30, (31, 1, 1601)),
        (365, (1, 1, 1602)),
        (365+31, (1, 2, 1602)),
        (365*3+31+28, (29, 2, 1604)),
        (365*3+31+29+2, (3, 3, 1604)),
        (5591, (23, 4, 1616)),
        (134326, (10, 10, 1968)),
        (146097, (1, 1, 2001)),
        (148253, (27, 11, 2006))
    ]
    for val, ex_res in tests:
        res = gregorian_date(val)
        print(f'{val} {res}...', end='')
        assert ex_res == res, f'the expected value is {ex_res}'
        print('OK')


        
        
        
test_gregorian_date()


```

Output:
```
0 (1, 1, 1601)...OK
30 (31, 1, 1601)...OK
365 (1, 1, 1602)...OK
396 (1, 2, 1602)...OK
1154 (29, 2, 1604)...OK
1157 (3, 3, 1604)...OK
5591 (23, 4, 1616)...OK
134326 (10, 10, 1968)...OK
146097 (1, 1, 2001)...OK
148253 (27, 11, 2006)...OK
```

### Día de la semana

Sabiendo que el 1/1/1601 fué lunes, podemos saber en qué dia de la semana cae cualquier fecha de las que estamos considerando (posteriores a 1/1/1601). Haz una función que calcule el día de la semana de una fecha dada. La función devolverá un entero: el lunes es el día 0 y el domingo el día 6. 

<!-- grade_id: cell-d4cd129e5a3af529, points: 0, grade: false, solution: true -->
```python
### 1 punto


def weekday(day, month, year):
    ### BEGIN SOLUTION
    assert is_correct(day, month, year), f'The date {day}-{month}-{year} is not correct'
    ndays = count_days_from_init(day, month, year)
    return ndays % 7
    ### END SOLUTION

```

```python
### Escribe aquí tus pruebas
```

<!-- grade_id: cell-85db837e377cd27d, points: 1, grade: true, solution: false -->
```python
def test_weekday():
    # n = count_days_from_init(3,10,2022)
    # for i in range(14):
    #     date = gregorian_date(n+i)
    #     wday = weekday(*date)
    #     print(f'({date}, {wday}),')
    tests = [
        ((3, 10, 2022), 0),
        ((4, 10, 2022), 1),
        ((5, 10, 2022), 2),
        ((6, 10, 2022), 3),
        ((7, 10, 2022), 4),
        ((8, 10, 2022), 5),
        ((9, 10, 2022), 6),
        ((10, 10, 2022), 0),
        ((11, 10, 2022), 1),
        ((12, 10, 2022), 2),
        ((13, 10, 2022), 3),
        ((14, 10, 2022), 4),
        ((15, 10, 2022), 5),
        ((16, 10, 2022), 6)
    ]

    for val, ex_res in tests:
        res = weekday(*val)
        print(f'{val} {res}...', end='')
        assert ex_res == res, f'the expected value is {ex_res}'
        print('OK')

test_weekday()
```

Output:
```
(3, 10, 2022) 0...OK
(4, 10, 2022) 1...OK
(5, 10, 2022) 2...OK
(6, 10, 2022) 3...OK
(7, 10, 2022) 4...OK
(8, 10, 2022) 5...OK
(9, 10, 2022) 6...OK
(10, 10, 2022) 0...OK
(11, 10, 2022) 1...OK
(12, 10, 2022) 2...OK
(13, 10, 2022) 3...OK
(14, 10, 2022) 4...OK
(15, 10, 2022) 5...OK
(16, 10, 2022) 6...OK
```