```python

```

# Suma Marciana

Se acabqa de publicar un hallazgo del rover perseverance en Marte. Se trata de una roca con la inscripcción:
      $$\begin{array}{ccc}
          \clubsuit   &   \diamondsuit       &    \spadesuit     \\
                      &  \clubsuit   &    \heartsuit \\
          \hline
          \diamondsuit        &  \spadesuit      &    \clubsuit
       \end{array}$$
Los científicos creen que se trata de una suma. Realiza un programa que calcule una lista de tuplas con todas las posibles soluciones a este enigma.
Puesto que la base numérica en la que está escrita también es desconocida, la función tener este dato como parámetro. La función debe llamarse `martian_sum`. Cada tupla debe ser 
de de la forma $(a_0, a_1, a_2, a_3)$ donde cada $0\leq a_k\leq b$, siendo $b$ la base y $a_0$ es el dígito de $\clubsuit$, $a_1$ es el valor de $\diamondsuit$, $a_2$ el valor de $\spadesuit$ y $a_3$ el valor de $\heartsuit$.
Por ejemplo si la base es 2, las soluciones son `[(0, 0, 0, 0), (0, 0, 1, 1)]`


<!-- grade_id: cell-a86ed308b1e232ed, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def martian_sum(base: int) -> list[int]:
    solutions = []
    club = 0
    while club < base:
        diamond = 0
        while diamond < base:
            spade = 0
            while spade < base:
                heart = 0
                while heart < base:
                    sum1 = base**2*club + base*diamond + spade
                    sum2 = base*club + heart
                    result = base**2 * diamond + base * spade + club
                    if sum1+sum2 == result:
                        solutions.append((club, diamond, spade, heart))
                    heart += 1
                spade += 1
            diamond += 1
        club += 1
    return solutions
### END SOLUTION
```

```python
martian_sum(4)
```

<!-- grade_id: cell-1c5b94aec65d7c82, points: 1, grade: true, solution: false -->
```python
from testing import runtests

def test_martian_sum():
    def wrapper(base):
        lst = martian_sum(base)
        lst.sort()
        return lst
    
    tests = [
        ((2,), [(0, 0, 0, 0), (0, 0, 1, 1)]), 
        ((4,), [(0, 0, 0, 0), (0, 0, 1, 3), (1, 1, 3, 2), (2, 3, 1, 1)]),
        ((10,), [(0, 0, 0, 0),
                 (0, 0, 1, 9),
                 (1, 1, 3, 8),
                 (2, 2, 5, 7),
                 (3, 3, 7, 6),
                 (4, 4, 9, 5),
                 (5, 6, 1, 4),
                 (6, 7, 3, 3),
                 (7, 8, 5, 2),
                 (8, 9, 7, 1)])
    ]
    
    runtests(tests, wrapper)
test_martian_sum()
```

Las soluciones para la base 4, son las siguientes: `[(0, 0, 0, 0), (0, 0, 1, 3), (1, 1, 3, 2), (2, 3, 1, 1)]`
Si asumimos la hipótesis que todos los valores deben ser distintos, ninguna de las soluciones anteriores es válida. 
Para ver si existe una solución es válida, primero vamos a realizar una función que dada una lista, nos inique si todos los elementos son distintos. La función se debe llamar `are_all_different`

<!-- grade_id: cell-255cf5fee91e2f70, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def all_different(lst: list[int]) -> bool:
    i = 1
    enc = False
    while i<len(lst) and not enc:
        j = 0
        while j<i and lst[i]!=lst[j]:
            j = j + 1
        enc = j<i
        i = i + 1
    return not enc
### END SOLUTION
```

```python

```

<!-- grade_id: cell-8754a6ef72533f5e, points: 1, grade: true, solution: false -->
```python
def test_all_different():
    tests = [
        (([],), True),
        (([1],), True),
        (([1,2],), True),
        (([1,2,4,5,6,7,8,9],), True),
        (([1,1,2,4,5,6,7,8,9],), False),
        (([1,2,4,1,5,6,7,8,9],), False),
        (([1,2,4,5,6,7,8,9,1],), False),
    ]
    runtests(tests, all_different)
test_all_different()
```

```python

```

Ahora podemos intentar hacer un programa que nos indique la menor base para la que existe una solución.
La función se debe llamar `lower_base_martian` . Debe devolver la base y la tupla con una solución

<!-- grade_id: cell-3526dba2a41a1188, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def lower_base_martian() -> int:
    base = 2
    enc = False
    while not enc:
        sols = martian_sum(base)
        pos = 0
        while pos<len(sols) and not all_different(sols[pos]):
            pos = pos+1
        if pos<len(sols):
            enc = True
        else:
            base = base+1
    return base, sols[pos]
### END SOLUTION        
```

```python

```

<!-- grade_id: cell-d7b1c993db4107b7, points: 1, grade: true, solution: false -->
```python
def test_lower_base_martian():
    base, sol = lower_base_martian()
    assert base == 5 and sol == (3, 4, 2, 1), 'Error'
    print('OK')
test_lower_base_martian()
```