El siguiente programa va a pensar un número entre 0 y un número dado. No es necesario que lo entiendas.

```python
import random

class Secret(object):
    def __init__(self, n):
        self.max = n
        self.secret = random.randint(0,n+1)
        
    def get_max(self):
        return self.max
    
    def query(self, guess):
        if self.secret<guess:
            ans = -1
        elif self.secret>guess:
            ans = 1
        else:
            ans = 0
        return ans
```

Para decir que piense un número hay que decir lo siguiente

```python
secret = Secret(5)
```

Ahora podemos preguntar a `secret` sobre el número que creemos que es mediante el método `query` que devuelve -1 si el que ha adividado es más pequeño que el que pasamos, 1 si es mayor y 0 si lo hemos adivinado. 

```python
secret.query(5)
```

Output:
```
-1
```

```python
secret.query(3)
```

Output:
```
-1
```

```python
secret.query(2)
```

Output:
```
0
```

```python
Ahora puedes intentarlo otra vez
```

```python
secret = Secret(10)
```

```python
secret.query(0)
```

Output:
```
1
```

Realiza un programa que encuentre el número secreto de forma secuencial. Es decir, empezando en 0, luego 1, y así sucesivamente hasta que lo encuentre. El programa se debe llamar `find_secret_seq` y como parámetro objeto de la clase `Secret`. El programa debe devolver el número secreto.

<!-- grade_id: cell-b5b0d9e6018db888, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def find_secret_seq(secret: Secret) -> int:
    i = 0
    while secret.query(i)>0:
        i = i+1
    return i
### END SOLUTION
```

```python

```

<!-- grade_id: cell-b2a2c9678715d093, points: 1, grade: true, solution: false -->
```python
def test_find_secret_seq():
    secret = Secret(100)
    n = find_secret_seq(secret)
    assert n == secret.secret, 'Error'
    print('OK')
test_find_secret_seq()
```

Output:
```
OK
```

Prueba ese programa con un número grande, déjalo funcionar, vete a tomar algo y a lo mejor ha acabado

```python
s = Secret(10**10)
print(find_secret_seq(s))
```

```python

```

Intenta adivinar el siguiente, 

```python
s = Secret(1000)
```

```python
s.query(...)
```

Piensa en la estrategia e intenta programarla. Haz una función llamada `find_secret` que implemenete esa estrategia para adivinar el número. La función admite tiene como parámetro un objeto de clase `Secret`. Puedes usar el método `get_max` para saber el límite superior en el que se encuentra el número.

```python
s.get_max()
```

Output:
```
1000
```

<!-- grade_id: cell-7ad8e673b39cfd05, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def find_secret(s: Secret) -> int:
    li = 0
    ls = s.get_max()
    enc = False
    while not enc:
        med = (li+ls)//2
        if s.query(med)>0:
            li = med
        elif s.query(med)<0:
            ls = med
        else:
            enc = True
            ans = med
    return ans
### END SOLUTION
```

<!-- grade_id: cell-d04e58ec395fb323, points: 1, grade: true, solution: false -->
```python
def test_find_secret():
    secret = Secret(10**10)
    n = find_secret(secret)
    assert n == secret.secret, 'Error'
    print('OK')
test_find_secret()
```

Output:
```
OK
```

Podemos probar ahora qué pasa cuando le pasamos un número grande.

```python
s = Secret(10**10)
print(find_secret(s))
```

O incluso números mayores

```python
s = Secret(10**100)
print(find_secret(s))
```